# __Programming Assignment 6a__


## Assignment Instructions 

Write three functions as described below. Be sure to test your code against the test cases! Submit your code once your tests are successfully passing. __Variable types are important in this one__ so be sure to read the directions closely below.

### __Function 1 :__


Write a function that takes one argument that is a string, and returns a list that contains every other character entered (starting from the character at the 0th position in the list).

##### Function 1 : Hints(click the arrow on the left to access)

- You can use the list() function to convert the string to a list of characters.
- You can use the slice operator to create a new list that only contains every other character in the original list. (Note that spaces will be treated as letters.)

### __Function 2 :__


Write a function that takes one argument that is a string, and returns an integer that represents the number of times the letter ‘e’ is found in the string.

### __Function 3 :__

Write a function that takes a string, then returns a tuple that contains only the first 5 characters in the string.

### __Main Program__

Now put all these three functions together to create a program that prompts a user to enter a sentence . The sentence is then split into individual letters using the first function and every other letter of the sentence is printed. Then,using the second function, the number of times the letter ‘e’  is found is counted and returned, followed by a function that returns a tuple that contains the first 5 characters of the string.

### Let's start coding ! 

```python
#install the necessary dependencies
import nose
from nose.tools import assert_equal
```

### Function 1 :

<!-- grade_id: cell-5771e291a6271b36, points: 0, grade: false, solution: true -->
```python
def every_other_letter(initial_string):
    ''' return a list,named initial_list, that contains every other letter of the initial_string,
        starting from the 0th position in the list '''
    my_list = list(initial_string)
    return my_list[::2]
every_other_letter('This is sentence 1')
```

Output:
```
['T', 'i', ' ', 's', 's', 'n', 'e', 'c', ' ']
```

```python
# Ungraded Tests - Test your function here
initial_string = "This is sentence 1"
every_other_letter(initial_string)
```

Output:
```
['T', 'i', ' ', 's', 's', 'n', 'e', 'c', ' ']
```

<!-- grade_id: cell-da43c75330edb865, points: 2, grade: true, solution: false -->
```python
# GRADED TEST CASE : Test your solution with the tests added in this cell
assert every_other_letter("This is sentence 1") == ['T', 'i', ' ', 's', 's', 'n', 'e', 'c', ' '], "Check your function to convert a string into a list of characters"
assert every_other_letter("Here is another sentence.") == ['H', 'r', ' ', 's', 'a', 'o', 'h', 'r', 's', 'n', 'e', 'c', '.'], "Check your function to convert a string into a list of characters"
```

### Function 2 : 

<!-- grade_id: cell-607fd14953620263, points: 0, grade: false, solution: true -->
```python
def count_letter(sentence_input):
    my_list = list(sentence_input)
    count = 0
    for i in my_list:
        if i == 'e':
            count += 1
    return count

count_letter('ee')
```

Output:
```
2
```

```python
# Ungraded ests - Test your function here

sentence_input = "I like to code" ## 2 e's
sentence_input2 = "I like coding" ## 0 e's
count_letter(sentence_input)
```

<!-- grade_id: cell-2ac44df2f08511b8, points: 3, grade: true, solution: false -->
```python
# GRADED TEST CASE : Test your solution with the tests added in this cell
assert count_letter("This is sentence 1") == 3 , "Soemthing seems to be wrong! Check your count_letter_e function  again "
assert count_letter("Here is a sentence with multiple e's") == 7, "Soemthing seems to be wrong! Check your count_letter_e function  again "
assert count_letter("This sentence has five e''s") == 5, "Soemthing seems to be wrong! Check your count_letter_e function  again "
```

### Function 3 : 

<!-- grade_id: cell-681ddc4dfa1a41aa, points: 0, grade: false, solution: true -->
```python
def tuple_first_five(word_str):
    ''' first, convert the string to a tuple, then return the first 5 letters of the tuple. Name the variable for tuple : five_tuple '''
    my_tuple = tuple(word_str)
    return my_tuple[:5]

tuple_first_five('halefcaetano')
```

Output:
```
('h', 'a', 'l', 'e', 'f')
```

```python
# Ungraded Tests - Test your function here
word_str = 'the quick brown fox jumps over the lazy dog'
tuple_first_five(word_str)
```

<!-- grade_id: cell-fbcef9f007d59468, points: 1, grade: true, solution: false -->
```python
# GRADED TEST CASE : Test your solution with the tests added in this cell
assert tuple_first_five("the quick brown fox jumps over the lazy dog") == ('t', 'h', 'e', ' ', 'q'), "Something is wrong! Check your function above." 
```

### __Putting it all Together in the Main Program__

Now, lets put all these functions together in one program!

```python
def every_other_letter(initial_string):
    ''' return a list,named initial_list, that contains every other letter of the initial_string,
        starting from the 0th position in the list '''
    my_list = list(initial_string)
    return my_list[::2]


def count_letter(sentence_input):
    my_list = list(sentence_input)
    count = 0
    for i in my_list:
        if i == 'e':
            count += 1
    return count


def tuple_first_five(word_str):
    ''' first, convert the string to a tuple, then return the first 5 letters of the tuple. Name the variable for tuple : five_tuple '''
    my_tuple = tuple(word_str)
    return my_tuple[:5]


def main():
    # prompt the user to enter a sentence
    user_string = input('Please enter a sentence: ')

    # Output every other letter using the function you created above
    print('Every other letter in the list is: ', every_other_letter(user_string))
    # Output the number of e's by calling the function you created above
    print('You entered', count_letter(user_string), "letter e's")
    # Output the first five letters as a tuple 
    print('The first five letters of the tuple are:',tuple_first_five(user_string))
    
#Do not remove this line as it is calling the main function
main()
```

Output:
```
Please enter a sentence:  halef
Every other letter in the list is:  ['h', 'l', 'f']
You entered 1 letter e's
The first five letters of the tuple are: ('h', 'a', 'l', 'e', 'f')
```

## Final Step: Submit your Work for Grading 

Great job on your work this week! To receive a final score for your work, please select the "Submit Assignment" button at the top of your lab.