# Programming Assignment 11b

## Objectives 

**The purpose of this assignment is to:**

- Revise how to open a file 

- Examine the data contained in a webpage

- Use the re package to extract interesting information, using **re.search() .group() and .findall()**

## Assignment Instructions

In this week's assignment, we will use regular expressions to extract data. Check out this [webpage](https://www.bsu.edu/academics/collegesanddepartments/online/academic-programs/masters/ms-data-science) about the Ball State Data Science Master's program.  Notice that there are a few phone numbers here and there.  There are also zip codes on this page. We will be extracting this data. 

For this purpose, we have added a file named **MSDSwebpage.txt** in this lab. It is an HTML file for the webpage, but the extension has been changed to ".txt" to make it easier for you to view it on your computer using a program like Notepad or TextEdit.

### Function 1: process_webpage(filename)

Write a function that accepts the name of a file as a string, opens the file, and returns the entire contents of the file as a string.  

Your functions should include error handling, and if the file does not exist in the directory (for example, named “doesnotexist.txt”), it should return the string “There is no file named doesnotexist.txt in this directory.”


##### Function 1 : Hints

Check the programming assignment for Module 8.  You’ll need a similar structure.

### Function 2: get_phone_numbers(web_string)

Write a function that accepts a string named **web_string**, which contains the contents previously extracted from the file by function 1.  This function should then extract phone numbers and return a list of these numbers.

##### Function 2 : Hints

- There are several ways to solve this problem.  One way is to use split() to split the input string into separate lines (use ‘\n’ with split), and then process them with a for loop.
- Be careful, as there may be more than one phone number found on a single line.  You may want to consider using re.findall() instead of re.search().

### Function 3: get_5digit_nums(web_string)

Write a function that accepts a string named **web_string**, which contains the contents previously extracted from the file by function 1.  This function should then extract all sets of 5 digit numbers (similar to zip codes) that are found in the file, returning a list of these numbers.


##### Function 3 : Hints

This function will be very similar to function 2.  Note, however, that what is found are not actually zip codes, but instead other numbers that are found in the file.

### Main Program 

For main program, we put all the 3 functions together. The code has been provided for you. 
Feel free to modify the code in main to further your understanding, but be sure that the three functions meet the requirements listed above.

**Sample output**

This file has the following items which may be phone numbers:
['800-872-0369', '765-285-1581', '800-382-8540', '765-289-1241']


This file has the following items which may be zip codes:
['00000', '00027', '96131', '16889', '11509', '16889', '11509', '00000', '00027', '00000', '00027', '15000', '63819', '31467', '60000', '63819', '31467', '60000', '63824', '17040', '20000', '63819', '31467', '40000', '63819', '31467', '60000', '63819', '31467', '60000', '63819', '31467', '60000', '63819', '31467', '60000', '11231', '66518', '12688', '76216', '23035', '40962', '54040', '16910', '72141', '49343', '63781', '44737', '46327', '49337', '49337', '49593', '00000', '00027', '35444', '11390', '36676', '35444', '09228', '81842', '22913', '66779', '66779', '66779', '49098', '66779', '66779', '70777', '43222', '64935', '62174', '47014', '70777', '43222', '64935', '62174', '47014', '62174', '47014', '62174', '47014', '62174', '47014', '54617', '68547', '57145', '70777', '43222', '64935', '54617', '68547', '57145', '54617', '68547', '57145', '54617', '68547', '57145', '54617', '68547', '57145', '90600', '70777', '43222', '64935', '90600', '90600', '90600', '90600', '95714', '28846', '70777', '43222', '64935', '95714', '28846', '95714', '28846', '95714', '28846', '95714', '28846', '77887', '80123', '49186', '70777', '43222', '64935', '77887', '80123', '49186', '77887', '80123', '49186', '77887', '80123', '49186', '77887', '80123', '49186', '63819', '31467', '60000', '47306', '47306']


## Let's Start Coding!

### Using this Notebook

Notice that some of the cells are not able to be edited – they are just meant to be run to test your code, and perform the auto-grading. You should begin at the top and execute each code cell, adding your code when you see # YOUR CODE HERE. If you have written the functions correctly, you should see no errors. 

When you execute the TEST CASES cell, what you want to see is no output. You may want to intentionally make the function give a wrong result, so that you can see what the error message looks like – and then fix your function. Once you execute the TEST CASES cell, and get no output, then your function is most likely correct! 

```python
# Install required dependencies 
import re, inspect

```

#### Function 1: process_webpage(filename)

<!-- grade_id: cell-9e602afcafd65048, points: 0, grade: false, solution: true -->
```python
def process_webpage(file_name):
    ''' open file named filename and read all contents into a single string, 
        returning the string '''
    try:
        with open(file_name, 'r') as datafile:
            content = ''
            line_list = datafile.readlines()
            for i in line_list:
                content += i
        return content
    except FileNotFoundError:
        return f"There is no file named {file_name} in this directory."
        
```

####  Function 2: get_phone_numbers(web_string)

<!-- grade_id: cell-9828bc958fbc5bd1, points: 0, grade: false, solution: true -->
```python
def get_phone_numbers(web_string):
    ''' processed the web_string and extract phone numbers, returning a list of numbers '''
    contents = ''
    for i in web_string:
        contents += i 
    r = re.compile(r'\d\d\d-\d\d\d-\d\d\d\d')
    results = r.findall(contents)
    return results
```

#### Function 3: get_5digit_nums(web_string)

<!-- grade_id: cell-7d10da9d6bf30ae0, points: 0, grade: false, solution: true -->
```python
def get_5digit_nums(web_string): 
    ''' take the string and extract 5 digit numbers, similar to zipcodes  '''
    contents = ''
    for i in web_string:
        contents += i 
    r = re.compile(r'\d\d\d\d\d')
    results = r.findall(contents)
    return results
```

## Testing your code with the Main program

The main program has been provided for you, to use in testing. Feel free to modify the code in main to further your understanding, but be sure that the three functions meet the requirements listed above.

Note : Do not change the variable names in the main() function below as they can impact the graded test cases.

```python
def main():
    main.web_text = process_webpage('MSDSwebsite.txt')
    main.phone_numbers = get_phone_numbers(main.web_text)
    main.fivedigit_nums = get_5digit_nums(main.web_text)
    print('This file has the following items which may be phone numbers:')
    print(main.phone_numbers)
    print('This file has the following items which may be zip codes:')
    print(main.fivedigit_nums)
    
    #returning the first 50 characters of the content
    print("\nReturning first 50 characters of the file :")
    return(main.web_text[0:50])

main()


```

Output:
```
This file has the following items which may be phone numbers:
['800-872-0369', '765-285-1581', '800-382-8540', '765-289-1241']
This file has the following items which may be zip codes:
['00000', '00027', '96131', '16889', '11509', '16889', '11509', '00000', '00027', '00000', '00027', '15000', '63819', '31467', '60000', '63819', '31467', '60000', '63824', '17040', '20000', '63819', '31467', '40000', '63819', '31467', '60000', '63819', '31467', '60000', '63819', '31467', '60000', '63819', '31467', '60000', '11231', '66518', '12688', '76216', '23035', '40962', '54040', '16910', '72141', '49343', '63781', '44737', '46327', '49337', '49337', '49593', '00000', '00027', '35444', '11390', '36676', '35444', '09228', '81842', '22913', '66779', '66779', '66779', '49098', '66779', '66779', '70777', '43222', '64935', '62174', '47014', '70777', '43222', '64935', '62174', '47014', '62174', '47014', '62174', '47014', '62174', '47014', '54617', '68547', '57145', '70777', '43222', '64935', '54617', '685
```

## Graded Test Cases 

The next cells shows how the auto-grading tests are executed. Several possible cases are run through the function, and the results are tested against that actual correct answer.
**If you run the test cell and there is no output, then the tests all passed – Great!**

<!-- grade_id: cell-ae95d54b3be13d1a, points: 4, grade: true, solution: false -->
```python
#GRADED TEST CASE 1a : Checking if the Correct file is processed 
errorCheck = False
error_statement = "There is no file"

#calling the main function to get the output
results = main()
if error_statement in results:
    errorCheck = False
else:
    errorCheck = True

assert errorCheck, "It seems like you have used an incorrect file. Please check if you have processed the correct file! " 
```

Output:
```
This file has the following items which may be phone numbers:
['800-872-0369', '765-285-1581', '800-382-8540', '765-289-1241']
This file has the following items which may be zip codes:
['00000', '00027', '96131', '16889', '11509', '16889', '11509', '00000', '00027', '00000', '00027', '15000', '63819', '31467', '60000', '63819', '31467', '60000', '63824', '17040', '20000', '63819', '31467', '40000', '63819', '31467', '60000', '63819', '31467', '60000', '63819', '31467', '60000', '63819', '31467', '60000', '11231', '66518', '12688', '76216', '23035', '40962', '54040', '16910', '72141', '49343', '63781', '44737', '46327', '49337', '49337', '49593', '00000', '00027', '35444', '11390', '36676', '35444', '09228', '81842', '22913', '66779', '66779', '66779', '49098', '66779', '66779', '70777', '43222', '64935', '62174', '47014', '70777', '43222', '64935', '62174', '47014', '62174', '47014', '62174', '47014', '62174', '47014', '54617', '68547', '57145', '70777', '43222', '64935', '54617', '685
```

<!-- grade_id: cell-b5f6d9af3526cd08, points: 1, grade: true, solution: false -->
```python
#GRADED TEST CASE 1b : Checking if an incorrect file is caught or processed
#In this cell , we are checking whether the function raises an error when an incorrect file is passed to it
error_check = bool
error_statement = "There is no file"
file_name = "SomeFile.txt"
results = process_webpage(file_name)

if error_statement in results:
    error_check = True
else :
    error_check = False


assert error_check, " Please check if you have processed the file correctly! " 

```

<!-- grade_id: cell-3fdc9f0409dddf9b, points: 5, grade: true, solution: false -->
```python
#GRADED TEST CASE 2 : Checking if a list of numbers is extracted from the file

error_check = False
result_list = ['800-872-0369', '765-285-1581', '800-382-8540', '765-289-1241']

if result_list == main.phone_numbers :
    error_check = True
else :
    error_check = False
    
assert error_check, "Please check your function and make sure you are processing the right file"
```

<!-- grade_id: cell-16e3c2b19214f119, points: 5, grade: true, solution: false -->
```python
#GRADED TEST CASE 3 : Checking if only 5-digit numberic data is extracted from the file

error_check_regex = bool
error_check_len = bool 

#using a list of correct 5-digit numbers similar to zipcodes
expected_pincodes = ['00000' , '76216', '22913', '66779', '47306']

#using a list of incorrect outputs to test 
incorrect_pincodes = [ 'B2345' , '555555' ,'200' ,'XB345']

function_result = main.fivedigit_nums

check = all(items in function_result for items in incorrect_pincodes)

#Test for checking the regex used in the function
if check is True :
    error_check_regex = False
else :
    error_check_regex = True
    
#Test case to check the length of each extracted number
for each_zipcode in function_result :
    if len(each_zipcode) != 5 :
        error_check_len = False
    else :
        error_check_len = True
        

assert error_check_regex , "Please check your function. Make sure you are extracting from the right file and only numeric codes are extracted"
assert error_check_len , "Please check your function. Make sure you are extracting from the right file and only 5-digit numbers are extracted"
```

## Submit your Work for Grading

Congratulations on completing this assignment.

To receive a final score for your work, please select the "Submit Assignment" button at the top of your lab.