# __Programming Assignment 5__

## __Assignment Overview__

Write a program that takes user input of 5 integers , stores them in a list and then loops through the list to calculate the average of those 5 numbers as the output.Additional intructions are added in the comments below, make sure to read through them carefully.

Be sure to test your code against the test cases! Submit your code once your tests are successfully passing.

### Lets start coding!! 

First, you will need to ask the user to enter 5 numbers and store them in a list . [__Hint :__ Try reading the input from user in a loop to store in a list] 

<!-- grade_id: cell-7a82fef6d1210fe2, points: 0, grade: false, solution: true -->
```python
l = []
s = 0
for i in range(5):
    l += [int(input('Enter a homework grade: '))]
    
for i in l:
    s += i

total = s / len(l)
print('Your average homework grade is:', total)

```

Output:
```
Enter a homework grade:  7
Enter a homework grade:  7
Enter a homework grade:  7
Enter a homework grade:  7
Enter a homework grade:  7
Your average homework grade is: 7.0
```

```python
#print out the list of numbers 

print("You entered the following homework grades: ", homework_grades)
```

### Calculating the average 

<!-- grade_id: cell-3bd0d4a350efc12f, points: 0, grade: false, solution: true -->
```python
# Use a for loop (with range()) to add up all of the grades in the list. 
# We have created a function that takes the list of numbers as parameter 
# and contains a variable named average already. Add your code in the mentioned spaces

def calc_average(grade_list) :

    average = 0
    
    # REMEMBER to remove the error return statement before running the cell and submitting the assignment
    # YOUR CODE HERE
    raise NotImplementedError()

    ## After the for loop finishes, divide the total by 5, in order to calculate the average.
    # YOUR CODE HERE
    raise NotImplementedError()
    
    return average

```

### Testing your work 

```python
# UNGRADED TEST CASES : You can use your own test lists to validate your code 
def calc_average(l):
    s = 0
    for i in l:
        s += i

    total = s / len(l)
    return total


your_list = [0,4,7,19,5]
calc_average(your_list)

```

Output:
```
7.0
```

<!-- grade_id: cell-49764892bbc6c937, points: 5, grade: true, solution: false -->
```python
# GRADED TEST CASES 1: You will not be able to edit this cell, but can run it for testing
'''Test your solution for autograding with the tests added below in this cell'''
### BEGIN TESTS
assert calc_average([1,3,4,5,6]) == 3.8 , "Please check your code for calculating the average ." 
assert calc_average([0,4,7,19,5]) == 7 , "Please check your code for calculating the average ."
### END TESTS
```

<!-- grade_id: cell-456f37f799c7c183, points: 5, grade: true, solution: false -->
```python
# GRADED TEST CASES 2: You will not be able to edit this cell, but can run it for testing

import inspect

# Check - for loop implementation
for_loop_found = False

calc_avg_lines = inspect.getsource(calc_average)
code_lines = calc_avg_lines.splitlines()

for line in code_lines:
    if "for" in line and "in" in line and ":" in line:
        for_loop_found = True
        
### BEGIN TESTS
assert for_loop_found == True, "Please verify you've implemented a for loop in calc_average."
### END TESTS
```

### Final Step :  Submit your Work for Grading

Great job on your work this week! To receive a final score for your work, please select the "Submit Assignment" button at the top of your lab.
