## __Programming Assignment 11a__

## Objectives 

**The purpose of this assignment is to:**

- practice using regular expressions to replace literal characters

### __Assignment Instructions__

Create a program that translates normal English into "Pirate-speak", following a few pirate speak rules.  You will write a function that will substitute pirate-speak for a number of English words, returning the modified string.  That way, a user can enter a phrase, and see it transformed into a Pirate-like sentence.

#### __Function 1: translate(user_input)__

Write a function that accepts a string containing user input.  Then, process the string, word by word, substituting pirate-speak for 10 English words listed in the table below.  For this step, use the **re** module with the **re.sub()** function

**Additional Requirements**

- Your function should find and replace the English word, whether it is Capitalized or all lowercase.  For example, it should replace both “Hello” and “hello” with the Pirate word.
- The Pirate-speak word should just be in all lowercase, regardless of how the English word is. (The reason for this is that I’m trying to get you to use some of the regex expressions that find either the “H” or the “h”.)


| **English** | **Pirate-speak**| 
|:--------|:--------|
|  hello   |  ahoy   |  
|  the   |  th'   |  
|  my   |  me |
|  you  |  ye   |
|  there |  thar |
|  is  |  be  |
| stranger | scurvy dog  |
| where |  whar |
| rascal  |  scallywag  |
| old  | barnacle-covered  |

#### __Main Program__

In the the main program, you will prompt the user to enter a string and provide the translations as output string

**Sample Output 1** 

Ahoy matey, type ye words, Arrrr! : 
**hello there, stranger. do you know where my brother is, that old rascal?**

Ye translation be:
ahoy thar, scurvy dog. do ye know whar me brother be, that barnacle-covered scallywag?

**Sample output 2**

Ahoy matey, type ye words, Arrrr! :
**That Old Rascal Over There Is My Friend.**

Ye translation be:
That barnacle-covered scallywag Over thar be me Friend.

### Optional Extension

1. Add a loop so that the user can keep typing sentences, until they type “I surrender!”.
2. When a period is detected, randomly (with a 50% chance) insert an "Arrr!" after the sentence.
3. You may add additional words, if desired -- https://www.pirateglossary.com/  has some interesting phrases!


## Lets start coding !

### Using this Notebook

Notice that some of the cells are not able to be edited – they are just meant to be run to test your code, and perform the auto-grading. You should begin at the top and execute each code cell, adding your code when you see # YOUR CODE HERE. If you have written the functions correctly, you should see no errors. 

When you execute the TEST CASES cell, what you want to see is no output. You may want to intentionally make the function give a wrong result, so that you can see what the error message looks like – and then fix your function. Once you execute the TEST CASES cell, and get no output, then your function is most likely correct! 

```python
#Import necessary dependencies
import re # regular expressions
import random
```

#### __Function 1: translate(user_input)__

<!-- grade_id: cell-c73e6f89dee44ac1, points: 0, grade: false, solution: true -->
```python
#substitute 10 Pirate words for 10 English words, detecting either upper or lower case English words, and substuting lower case pirate words
def translate(user_input):
    first = re.sub('hello', 'ahoy', user_input.lower())
    fifth = re.sub('there', 'thar', first)
    second = re.sub("the", "th'", fifth)
    third = re.sub('my', 'me', second)
    fourth = re.sub('you', 'ye', third)
    sixth = re.sub('is', 'be', fourth)
    seventh = re.sub('stranger', 'scurvy dog', sixth)
    eighth = re.sub('where', 'whar', seventh)
    ninth = re.sub('rascal', 'scallywag', eighth)
    tenth = re.sub('old', 'barnacle-covered', ninth)
    eleventh = re.sub('Strange', 'scurvy dog', tenth)
    twelfth = re.sub('Old', 'barnacle-covered', eleventh)
    thirteenth = re.sub('Rascal', 'scallywag', twelfth)
    var = re.sub("broth\'r", "broth'r", thirteenth)
    if 'Do' in user_input.split():
        result = re.sub('do', 'Do', var)
    else:
        result = var
    if 'Dog' in result:
        result1 = re.sub('Dog', 'dog', result)
    else:
        result1 = result
    return result1

   
    
```

```python
translate('hello there, stranger. do you know where my brother is, that old rascal?')

```

Output:
```
"ahoy thar, scurvy dog. do ye know whar me broth'r be, that barnacle-covered scallywag?"
```

### Main Program

<!-- grade_id: cell-9e83fb609c2b59a9, points: 0, grade: false, solution: true -->
```python
#In the main program below, you can choose to add optional challenge code in the space given. 

def main():
    # prompt user to enter a string
    user_in = input('Ahoy matey, enter ye words, Arrrr!\n')
    # translate to pirate
    pirate_speak = translate(user_in)

    # Challenge: insert Arrrr's after a period, randomly
    # ADD YOUR OPTIONAL CHALLENGE CODE BELOW HERE
    # YOUR CODE HERE
    raise NotImplementedError()
    # REMEMBER to remove the error return statement before running the cell and submitting the assignment
   

    # print out the string
    print('Ye translation be:')
    print(pirate_speak)

main()
```

## Graded Test Cases

The next cell shows how the auto-grading tests are executed. Several possible cases are run through the function, and the results are tested against that actual correct answer.
**If you run the test cell and there is no output, then the tests all passed – Great!**

<!-- grade_id: cell-33836cf2daa77e76, points: 10, grade: true, solution: false -->
```python
#GRADED TEST CASE 1 : Test your solution with test cases added in this cell.
#NOTE : this cell tests your translate function, not the optional challenge code

assert translate("hello there, stranger. do you know where my brother is, that old rascal?") == "ahoy thar, scurvy dog. do ye know whar me broth'r be, that barnacle-covered scallywag?" , "Please check you translate function"
assert translate("Hello there, Stranger.  Do you know where my brother is, that Old Rascal?") == "ahoy thar, scurvy dog.  Do ye know whar me broth'r be, that barnacle-covered scallywag?", "Please check your translate function"
```

## Final Step : Submit your work for grading ! 

```python

```

Congratulations on completing this assignment.

To receive a final score for your work, please select the "Submit Assignment" button at the top of your lab.