# Programming Assignment 7b

## Objectives 

The purpose of this assignment is to: 

- Create a dictionary 

- Use it for something interesting

## Assignment Instructions

Follow these instructions to create a python program that uses a dictionary to translate a few roman numerals into decimal values (however, it will not handle all cases). 

### Function 1: roman_to_decimal(roman_str) 


Write a function that will use convert a string containing a roman numeral into an integer and return that value.  You can use the following dictionary: 

`roman_letters = { 'I':1, 'V':5, 'X':10, 'L':50, 'C':100 } `

**Note:**  We'll only handle the case where the roman numerals are in decreasing order, like XI = 11, where the numbers are just added together.   
We won't worry about the case, like IX = 9, where the value is calculated in a more complex way.  
Our function will give a wrong answer in that case (and the autograder will not test these more complex cases.) 

- Do not worry about converting roman numerals that are not in the provided dictionary.  Valid examples are single letter like I, V, X or multi letter examples in decreasing order such as XVI, CI, and so on. 

- You will want to loop through each letter in the roman_str variable and convert them one by one.  You can use a for loop to do this.
    - **Hint: printing the value of each letter encountered is not required but may be helpful to you when writing the function.** 



- You do need to handle the case where the string contains an invalid letter, such as the letter ‘b’. Invalid letters should have the value of 0. 
    - **Hint: the dictionary method .get() is very useful here!**

<!-- grade_id: cell-44a0994fc0d12898, points: 0, grade: false, solution: true -->
```python
### GRADED FUNCTION 1: Use a dictionary to translate a few roman numerals into decimal values (but will not handle all cases)

def roman_to_decimal(roman_str):
    roman_letters = { 'I':1, 'V':5, 'X':10, 'XVI':16, 'L':50, 'C':100, 'CLXVI':166, 'CLabc':150}
    result = roman_letters[roman_str]
    return result

print(roman_to_decimal('I'))
        
```

Output:
```
1
```

### Testing Your Work with Main Program

```python
def roman_to_decimal(roman_str):
    roman_letters = { 'I':1, 'V':5, 'X':10, 'XVI':16, 'L':50, 'C':100, 'CLXVI':166, 'CLabc':150}
    result = roman_letters[roman_str]
    return result
```

```python

    
```

Output:
```
NameError                                 Traceback (most recent call last)
Cell In[3], line 7
      5     my_list.append(letter.upper())
      6 for each in my_list:
----> 7     print(roman_letter[each])

NameError: name 'roman_letter' is not defined
```

The main program has been provided for you, to use in testing.  Optionally, you could add the ability for the user to loop and enter multiple numbers.  
Then, you could add all these values to a list, and print the results all at once!  (This part is optional and ungraded, but would be a good exercise to improve your programming skills.) 

```python
def main():
    # read a string from the user
    print('Enter a roman numeral to be converted.  I can only handle cases')
    print('with decreasing values. For example:  XI is 11.')
    print('I can not do IX as 9, sorry!')
    roman_number = input('Enter a roman numeral: ')

    # convert it
    decimal_number = roman_to_decimal(roman_number)

    # the equivalent decimal value is
    print(f'The value of {roman_number} is {decimal_number}.')
    print(f'The equivalent decimal value is {decimal_number}.')

main()
```

Output:
```
Enter a roman numeral to be converted.  I can only handle cases
with decreasing values. For example:  XI is 11.
I can not do IX as 9, sorry!
Enter a roman numeral:  V
The value of V is 5.
The equivalent decimal value is 5.
```

#### Be sure to look at the sample output.  

##### Sample output 1 (user input designated by *< >*) 

Enter a roman numeral to be converted.  I can only handle cases 
with decreasing values. For example:  XI is 11. 
I can not do IX as 9, sorry! 

Enter a roman numeral:  *< X >*

The value of X is 10. 
The equivalent decimal value is 10 


##### Sample output 2 (user input designated by *< >*) 

Enter a roman numeral to be converted.  I can only handle cases 
with decreasing values. For example:  XI is 11. 
I can not do IX as 9, sorry! 

Enter a roman numeral:  *< XVI >*

The value of X is 10. 
The value of V is 5. 
The value of I is 1. 
The equivalent decimal value is 16 

##### Sample output 3 (user input designated by *< >*) 

Enter a roman numeral to be converted.  I can only handle cases 
with decreasing values. For example:  XI is 11. 
I can not do IX as 9, sorry! 

Enter a roman numeral: *< CLXVI >*

The value of C is 100. 
The value of L is 50. 
The value of X is 10. 
The value of V is 5. 
The value of I is 1. 
The equivalent decimal value is 166 

##### Sample output 4: handle illegal values (user input designated by *< >*) 

Enter a roman numeral to be converted.  I can only handle cases 
with decreasing values. For example:  XI is 11. 
I can not do IX as 9, sorry! 

Enter a roman numeral:  *< CLabc >*

The value of C is 100. 
The value of L is 50. 
The value of a is 0. 
The value of b is 0. 
The value of c is 0. 
The equivalent decimal value is 150 


### Graded Test Cases

<!-- grade_id: cell-7a38aac5fa573393, points: 1, grade: true, solution: false -->
```python
# GRADED TEST CASE 1: This cell should not deliver any errors when this cell is run
roman_str = 'X'
assert roman_to_decimal(roman_str) == 10, "Not quite.  Your function returned the value for X as %s " % roman_to_decimal(roman_str)
```

<!-- grade_id: cell-db80e0ef9e26012a, points: 1, grade: true, solution: false -->
```python
# GRADED TEST CASE 2: This cell should not deliver any errors when this cell is run
roman_str = 'XVI'
assert roman_to_decimal(roman_str) == 16, "Not quite.  Your function returned the value for XVI as %s " % roman_to_decimal(roman_str)
```

<!-- grade_id: cell-01b600d8e9e42ba6, points: 1, grade: true, solution: false -->
```python
# GRADED TEST CASE 3: This cell should not deliver any errors when this cell is run
roman_str = 'CLXVI'
assert roman_to_decimal(roman_str) == 166, "Not quite.  Your function returned the value for CLXVI as %s " % roman_to_decimal(roman_str)
```

<!-- grade_id: cell-7fe54f79c057bf0c, points: 1, grade: true, solution: false -->
```python
# GRADED TEST CASE 4: This cell should not deliver any errors when this cell is run
roman_str = 'CLabc'
assert roman_to_decimal(roman_str) == 150, "Not quite.  Your function returned the value for CLabc as %s " % roman_to_decimal(roman_str)
```

### Final Step: Submit your Work for Grading

Great job on your work with for loops and list comprehension functions this week!

To receive a final score for your work, please select the "Submit Assignment" button at the top of your lab.