# Programming Assignment 3b

## Assignment Instructions

Write two functions, described below. Be sure to test your code against the test cases! Submit your code once your tests are passing.  

*Note: It is possible to earn partial credit for passing some of the tests, but not passing others.*

### __Function 1__:  

Write a function to calculate the area of a circle, using the formula `pi*r*r`.  If you want, you can use the constant math.pi instead of entering 3.14159.  

### __Function 2__:  

Write a function to calculate the circumference of a circle, using the formula `2*pi*r`. If you want, you can use the constant math.pi again.


### Extension (Optional)  

If you want to turn these two functions into a full program, you could write more lines of code (refer to the live coding videos for how to structure the main parts of your program, and where to place these functions).  This part is totally optional! 


Follow these instructions to create a python program that will give the user a menu of 3 options to choose from, and then take a different action depending on which option they choose.  They can either calculate the area of a circle, the circumference of a circle, or quit the program.  The program will continue to ask them for these options, until they choose to quit. 

* Prompt the user to select from these options:  Enter 1 to calculate the area, 2 to calculate the circumference, 0 to exit the program. 

* If they type a 1, prompt the user to enter a radius as a floating point number, then call your function to calculate the area.  Then, loop back to step 2. 

*  If they type a 2, prompt the user to enter a radius as a floating point number, then call your function to calculate the circumference. Then, loop back to step 2. 

*  If they type a 0, exit the program. 

## Let's start coding!

#### Using this Notebook

Notice that some of the cells are not able to be edited – they are just meant to be run to test your code, and perform the auto grading. You should begin at the top and execute each code cell, adding your code when you see # YOUR CODE HERE. If you have written the functions correctly, you should see no errors.

When you execute the TEST CASES cell, what you want to see is no output. You may want to intentionally make the function give a wrong result, so that you can see what the error message looks like – and then fix your function. Once you execute the TEST CASES cell, and get no output, then your function is most likely correct!

```python
# Install Necessary Dependencies
import math, nose
from nose.tools import assert_almost_equal
```

### Function 1: area

<!-- grade_id: cell-f65fe616cbb0881f, points: 0, grade: false, solution: true -->
```python
### GRADED FUNCTION 1: Write a function to calculate the area of a circle.
def area(radius):
    ''' return the area of a circle of this radius '''
    pi = math.pi
    circle = pi * radius * radius
    return circle
    raise NotImplementedError()
number = int(input('Enter a  to calculate the area of the circle.'))
print(f'The area of the cirlce is {"%.2f" % area(number)}')
```

Output:
```
Enter a  to calculate the area of the circle. 20
The area of the cirlce is 1256.64
```

The cell below gives you an opportunity to test your code in a way similar to how the auto-grader will test it.  The first line of code allows you to replace the ‘…’ with a number of your choice.  When you run the cell, the second line will call the function that contains your solution.  The third line then prints the results, so that you can manually test it – if you put these numbers into a calculator, are they correct?

The next cell shows how the auto-grading tests are executed.  Several possible cases are run through the function, and the results are tested against that actual correct answer. If you run the test cell and there is no output, then the tests all passed – Great!

If you’re curious, the function named assert_almost_equal() is a way to test floating point numbers – if you remember from the lectures, these are always rounded at the end, and so we cannot say that they are exactly equal, but only that they are pretty close to the same number. 

```python
# Manually test `area`
test_radius_for_area = 20 # replace ... with an appropriate value of your choice for the radius variable
test_area_function = area(test_radius_for_area)
print(f"radius of circle: {test_radius_for_area}, calculated area: {test_area_function}")
```

Output:
```
radius of circle: 20, calculated area: 1256.6370614359173
```

<!-- grade_id: cell-3ada116f80055b02, points: 5, grade: true, solution: false -->
```python
# TEST CASES: This cell should not deliver any errors if all tests are passing.
assert_almost_equal(area(1), math.pi, places=2, msg="Check your area function.")
assert_almost_equal(area(0), 0, places=2, msg="Check your area function with radius 0.")
assert_almost_equal(area(2), 4*math.pi, places=2, msg="Check your area function with radius 2.")
assert_almost_equal(area(10), 100*math.pi, places=2, msg="Check your area function with radius 10.")
```

### Function 2: Circumference

<!-- grade_id: cell-97bec1841199b27e, points: 0, grade: false, solution: true -->
```python
def circumference(radius):
    ''' return the circumference of a circle of this radius '''
    pi = math.pi
    circle = pi * 2 * radius
    return circle
    raise NotImplementedError()
number = int(input('Enter a  to calculate the circumference of the circle.'))
print(f'The area of the cirlce is {"%.2f" % circumference(number)}')
```

Output:
```
Enter a  to calculate the circumference of the circle. 40
The area of the cirlce is 251.33
```

```python
# Manually test `circumference`
test_radius_for_circumference = 40 # replace ... with an appropriate value of your choice for the radius variable
test_area_function = circumference(test_radius_for_circumference)
print(f"radius of circle: {test_radius_for_circumference}, calculated area: {test_area_function}")
```

Output:
```
radius of circle: 40, calculated area: 251.32741228718345
```

<!-- grade_id: cell-0cead43d8902123f, points: 5, grade: true, solution: false -->
```python
# TEST CASES: This cell should not deliver any errors if all tests are passing.
assert_almost_equal(circumference(1), 2*math.pi, places=2, msg="Check your circumference function.")
assert_almost_equal(circumference(0), 0, places=2, msg="Check your circumference function with radius 0.")
assert_almost_equal(circumference(2), 4*math.pi, places=2, msg="Check your circumference function with radius 2.")
assert_almost_equal(circumference(10), 20*math.pi, places=2, msg="Check your circumference function with radius 10.")
```

## Putting it all together

Now let's put it all together! Below is a program that will prompt the user (you in this case!) for two measurement values. 
The program will then use these measurements to calculate an area and circumference of a circle. You can try this out below and watch how your functions come to life in a real-world application!

```python
def main():
    
    while True:
        # ask if they want an area or circumference
        answer = input('Enter 1 for the area\nEnter 2 for the circumference\nEnter 0 to exit\n')

        # do what the user says
        if answer == '1':  # area
            # prompt user for radius
            rad = float(input('Enter a radius: '))
            print(f'The area is {area(rad):.2f}')
        elif answer == '2': # circumference
            # prompt user for radius
            rad = float(input('Enter a radius: '))
            print(f'The circumference is {circumference(rad):.2f}')
        elif answer == '0': # quit
            break
        else: # they entered something else|
            print('I do not recognize that input.')
main()

```

Output:
```
Enter 1 for the area
Enter 2 for the circumference
Enter 0 to exit
 1
Enter a radius:  20
The area is 1256.64
Enter 1 for the area
Enter 2 for the circumference
Enter 0 to exit
 2
Enter a radius:  20
The circumference is 125.66
Enter 1 for the area
Enter 2 for the circumference
Enter 0 to exit
 0
```

```python

```