# Programming Assignment 12b

## Assignment Instructions

This program asks the user to enter a sentence, verifying that only one sentence was entered.   
Then, it performs a sentiment analysis on the sentence, using both the `TextBlob` default sentiment analyzer and the `NaiveBayesAnalyzer`.  

For the default sentiment analyzer, it will print out one of 5 emoji characters, in the range of 😀 🙂 😐 🙁 😟  where the sentiments are in the ranges: 

* 1.0 and 0.6 means 😀 

* 0.6 and 0.35 means 🙂 

* 0.35 and -0.35 means 😐  

* -0.35 and -0.6 means 🙁  

* -0.6 and -1.0 means 😟 

For the `NaiveBayesAnalyzer`, it will just print a happy face 😀 or a sad face 😟, for positive and negative sentiments, respectively. 

As usual, a simple main program is provided.  You will write the three functions called by main().  (Your grade for this assignment is based on the performance of these three functions.).   
Closely examine the sample output below to determine the functionally required. 

#### Sample Outputs


**Sample Output 1:**  
Please enter a single sentence: 

**User Input:** Today is a great day! 

Default sentiment analyzer reports:   
Polarity: 1.0 😀   
NaiveBayesAnalyzer reports:   
Polarity: 😀 positive   

**Sample Output 2:**  
Please enter a single sentence.  

**User Input:** Today is an ok day.   

Default sentiment analyzer reports:   
Polarity: 0.5 🙂   
NaiveBayesAnalyzer reports:   
Polarity: 😀 positive   

**Sample Output 3:**  
Please enter a single sentence. 

**User Input:**  Today is a day. 

Default sentiment analyzer reports:   
Polarity: 0.0 😐   
NaiveBayesAnalyzer reports:   
Polarity: 😀 positive   

**Sample Output 4:**    
Please enter a single sentence.   

**User Input:** Today is a lousy day.   

Default sentiment analyzer reports:   
Polarity: -0.5 🙁   
NaiveBayesAnalyzer reports:   
Polarity: 😟 negative   


**Sample Output 5:**  
Please enter a single sentence.   

**User Input:** Today was a terrible day.   

Default sentiment analyzer reports:   
Polarity: -1.0 😟   
NaiveBayesAnalyzer reports:   
Polarity: 😟 negative   


**Sample Output 6:**  
Please enter a single sentence.   

**User Input:** Today was a day. Tomorrow is another one.   
 
You entered more than one sentence. 

**Sample Output 7:**  
Please enter a single sentence.   

**User Input:** Today was an amazing day!  Hopefully, tomorrow will also be great. 
 
You entered more than one sentence. 


## Let's start coding!

#### Using this Notebook 

Notice that some of the cells are not able to be edited – they are just meant to be run to test your code, and perform the auto grading. You should begin at the top and execute each code cell, adding your code when you see # YOUR CODE HERE.  
If you have written the functions correctly, you should see no errors.

When you execute the TEST CASES cell, what you want to see is no output. You may want to intentionally make the function give a wrong result, so that you can see what the error message looks like – and then fix your function.   
Once you execute the TEST CASES cell, and get no output, then your function is most likely correct!

```python
import nltk
from textblob import TextBlob
from textblob.sentiments import NaiveBayesAnalyzer
```

### Step 1: Write function num_sentences()

<!-- grade_id: cell-44a0994fc0d12898, points: 0, grade: false, solution: true -->
```python
### GRADED FUNCTION: Write the function "num_sentences"

def num_sentences(text):
    blob = TextBlob(text)
    return len(blob.sentences)
```

```python
## Try out your function 
text = 'This is sentence one. This is sentence two. There should be three sentences'
print(num_sentences(text))
```

Output:
```
3
```

### Step 2: Write function report_sentiment(text)

<!-- grade_id: cell-58341e91ad0b7cba, points: 0, grade: false, solution: true -->
```python
### GRADED FUNCTION:  Write a second function called "report_sentiment"

def report_sentiment(text):
    ''' return a string reporting the TextBlob default sentiment, in emoji '''
    blob = TextBlob(text)
    polarity = blob.sentiment.polarity
    if polarity <= 1 or polarity > 0.6:
        return f'Polarity: {polarity} 😀'
    elif polarity <= 0.6 or polarity > 0.35:
        return f'Polarity: {polarity} 🙂'
    elif polarity <= 0.35 or polarity > -0.35:
        return f'Polarity: {polarity} 😐'
    elif polarity <= -0.35 or polarity > -0.6:
        return f'Polarity: {polarity} 🙁'
    elif polarity <= -0.6 or polarity >= -1.0:
        return f'Polarity: {polarity} 😟'
```

```python
## Try out your function 
text1 = 'Today is a lousy day.'
text2 = 'Textblob is amazingly simple to use. What great fun!'

print(report_sentiment(text1))
print(report_sentiment(text2))
```

Output:
```
Polarity: -0.5 😀
None
Polarity: 0.39166666666666666 😀
None
```

### Step 3: Write function report_NBA_sentiment(text)

<!-- grade_id: cell-b8df4cc0186a657c, points: 0, grade: false, solution: true -->
```python
### GRADED FUNCTION:  Write a third function called "report_NBA_sentiment"

def report_NBA_sentiment(text):
    ''' return a string reporting the TextBlob Naive Bayes Analyzer sentiment, and emoji '''
    blob = TextBlob(text, analyzer= NaiveBayesAnalyzer())
    if blob.sentiment.classification == 'pos':
        return f'Polarity: 😀 positive'
    else: 
        return f'Polarity: 😟 negative'
```

```python
## Try out your function
text1 = 'Today is a lousy day.'
text2 = 'Textblob is amazingly simple to use. What great fun!'
blob = TextBlob(text1, analyzer= NaiveBayesAnalyzer())
print(blob.sentiment)
print(report_NBA_sentiment(text1))
print(report_NBA_sentiment(text2))




```

Output:
```
Sentiment(classification='neg', p_pos=0.3431499740529319, p_neg=0.6568500259470683)
Polarity: 😟 negative
Polarity: 😀 positive
```

### Let's bring it all together in main()

The main function is provided for you below and is readonly. You do not need to edit this function.

```python
def main():
    # get input
    corpus_text = input('Please enter a single sentence.')
    # verify it is a single sentence, using TextBlob
    if num_sentences(corpus_text) != 1:
        print('You entered more than one sentence.')
    else:
        # print the sentiment using the default TextBlob analyzer
        print('Default sentiment analyzer reports:')
        print(report_sentiment(corpus_text))

        # print the sentiment using ths NBA 
        print('NaiveBayesAnalyzer reports:')
        print(report_NBA_sentiment(corpus_text))

main()
```

Output:
```
Please enter a single sentence. helo there
Cell In[12], line 16
     13         print('NaiveBayesAnalyzer reports:')
     14         print(report_NBA_sentiment(corpus_text))
---> 16 main()

Cell In[12], line 5, in main()
      3 corpus_text = input('Please enter a single sentence.')
      4 # verify it is a single sentence, using TextBlob
----> 5 if num_sentences(corpus_text) != 1:
      6     print('You entered more than one sentence.')
      7 else:
      8     # print the sentiment using the default TextBlob analyzer

NameError: name 'num_sentences' is not defined
```

### Step 3: Testing Your Work

Your work will be autograded based on your function outputs for the test cases below. The next cell shows how the auto-grading tests are executed.  

Several possible cases are run through the function, and the results are tested against that actual correct answer. If you run the test cell and there is no output, then the tests all passed – Great!

<!-- grade_id: cell-19240d2478ff1f89, points: 2, grade: true, solution: false -->
```python
# GRADED TEST CASES 1: This cell should not deliver any errors when this cell is run
# Function num_sentences(text): tests the return of an expected number of sentences for corpus text

text = "Today was a day. Tomorrow is another one."
assert num_sentences(text) == 2, "Please verify your function is calculating the correct number of sentences."
```

<!-- grade_id: cell-972625f7aab31e18, points: 4, grade: true, solution: false -->
```python
# GRADED TEST CASES 2: This cell should not deliver any errors when this cell is run
# Function report_sentiment(text): tests a string and emoji sentiment are returned.

text =  "Today was an amazing day! Hopefully, tomorrow will also be great."
exp_sentiment = "Polarity: 0.7750000000000001 😀"

assert report_sentiment(text) == exp_sentiment, "Please verify that your function delivers the appropriate sentiment polarity value and corresponding emoji."
```

<!-- grade_id: cell-3a25b6898afbf1a6, points: 4, grade: true, solution: false -->
```python
# GRADED TEST CASES 3: This cell should not deliver any errors when this cell is run
# Function report_NBA_sentiment(text): tests a string and emoji is returned, based on Naive Bayes Analyzer sentiment

text = "Today is a lousy day."
exp_sentiment = "Polarity: 😟 negative"

assert report_NBA_sentiment(text) == exp_sentiment, "Please verify that your function delivers the appropriate sentiment polarity value and corresponding emoji."
```

### Step 4: Submit your Work for Grading

To receive a final score for your work, please select the "Submit Assignment" button at the top of your lab.