# Programming Assignment 13


```python

```

## Assignment Instructions

The purpose of this assignment is to: 

* Learn how to perform webscraping (extracting content from a webpage in order to analyze the text) using the `BeautifulSoup` library 

* Use the `Textatistic` library to generate a readability score using the Dale-Chall score 

* Report the sentiment polarity and subjectivity using TextBlob 

**Note:**  We really should clean up the scraped text more -- if you print what is returned from the BeautifulSoup library, 
you will see that there is a bunch of other junk in there, which is likely messing up the readability score and the polarity. For example, I loaded a first-grade reading level book located here:  http://www.funbrainjr.com/stories/BrainyBlueberryAndTheBackpackBalloon/ 
and the Textatistic library reported a grade level of 13 (Freshman in college!).   Most likely, if we removed more of the "junk" from the webpage and just focused on the text, it would work better.  Data Munging is a critical part of any analysis! 

For this assignment, three websites have been selected, for articles on the same event by three different U.S. reporting agencies: 

* https://www.nytimes.com/live/2022/04/07/us/ketanji-brown-jackson-vote-scotus

* https://www.cbsnews.com/news/ketanji-brown-jackson-supreme-court-senate-confirmation-first-black-woman/    

* https://www.foxnews.com/politics/senate-confirmation-vote-ketanji-brown-jackson-supreme-court  

We will write a python program to connect to these websites, download the text by webscraping, do some basic data munging, and then perform some analysis.  As we will see, we’ll encounter a few issues – but these are very common and it’s important to see how complex a task like this can be to get really accurate answers.  Be sure to read the Discussion section below to see what these issues are. 

This assignment is a great place to use good programming practice – and to break repeated tasks into functions.  Main will be provided, and you will write several functions that each perform a task.   

##### **Function 1: extract_text(addr)**

Write a function that takes a web address, connects to the website using the requests package, extracts the text using the BeautifulSoup package, and performs some basic data munging to remove HTML tags, and then returns the cleaned text. We often call this action webscraping, since we are quickly “scraping off” the text from a webpage. 
 

##### **Hints**


* We have used the requests package previously in Project 2.  You can retrieve all the text on a webpage using the command: 
`response = requests.get(website_address_string)`

    * The BeautifulSoup library (https://www.crummy.com/software/BeautifulSoup/ ) contains library routines that are useful for webscraping.  

    * We will need to include the library:   
    `from bs4 import BeautifulSoup`

    * And then to make the soup object, we use the following line of code:   
    `soup = BeautifulSoup(response.content, 'html.parser')`

    * Note that, if you are doing this at home, you will need to follow the instructions in the documentation above to install both the bs4 package and the html.parser.  However, this Coursera lab environment already has those packages installed. 

* We can then remove HTML tags automatically (but not perfectly, see the discussion below) using the soup command:   
`text = soup.get_text(strip=True)`

##### **Function 2: readability(text)**

Write a function that takes a string of text (the cleaned results of the webpage) and uses a library called Textatistic (https://www.erinhengel.com/software/textatistic/) to return the Dale-Chall readability score (see https://en.wikipedia.org/wiki/Dale%E2%80%93Chall_readability_formula which is a measure of the reading difficulty).   


##### **Hints**

* We will need to import the library:   
`from textatistic import Textatistic` 

* We can gen}erate a number of measures using Textatistic using the following code:   
`readability = Textatistic(text)` 

* The Dale-Chall readability score is then accessed this way:   
`readability.dalechall_score` 

##### **Function 3: sentiment(text)**

Write a function that accepts the text as an argument, then uses TextBlob to calulate and return the sentiment polarity score.

##### **Hints**


We did this last week. 

##### **Function 4: print_sentences(text, num_sentences)**

Write a function that will print the first N sentences of a text, where N is given by the number stored in the num_sentences variable.    
Use the TextBlob package to automatically determine sentences.  This function just prints directly, and does not return a value. 

##### **Hints**


* First, convert the text to a TextBlob. 
* Then, you can use code similar to this to print: for sentence in blob.sentences: 
* You’ll want to count how many sentences were printed, and then exit the loop after the requested number have been printed. 
* See the discussion below, as this function yields surprising results. 

## Discussion 

This assignment illustrates the messy nature of working with real data.  If you examine the output of your program and the sample output below, you’ll notice a few things: 

* The actual data returned by a webpage can change over time – and what you receive may not be exactly what I received when I ran my code to generate the sample output. 

* In the output below – you'll notice that the first website did NOT actually return the text that you’d see if you browsed to the page normally.  It detected that something was different in how the program was automatically connecting and refused to send the text over.  There are ways to get around this, but this is a pretty common occurrence, so I’m glad that we got to see it in this assignment! 

* The second website has a jumbled first sentence (#0) – it's catching a bunch of links for other stories, which were combined into one “sentence”.  This is not really reflective of the article’s writing, and in a serious study, we would need to do a more thorough job of cleaning the data. 

* The second sentence (#1) is also jumbled, but the article actually starts in the middle of that sentence. 

* The remaining sentences are mostly ok – and the rest of the article follows. 

* The third website’s first 5 sentences are pretty much useless, if we are trying to measure the reading level or sentiment of the article. 

* Given all that, the sentiment and Dale-Chall scores are probably not very accurate!  We’d need to do a better job extracting the actual content of the article before doing the analysis.   

Be careful when doing analysis like this – and don’t just blindly trust the numbers returned.  Sometimes the data passed to the functions is not what you expect.  (There is the phrase “garbage in – garbage out”, which applies to many areas of both data analytics and computer science!) 

#### Sample Outputs

(Note that webpages are dynamic, and many include advertising and suggested further reading, which change over time.  Your values may be slightly different than those shown below.)

**Sample Output 1:**

Text from https://www.nytimes.com/live/2022/04/07/us/ketanji-brown-jackson-vote-scotus has results: 
Reability Dale-Chall score: 15.88 
Sentiment Polarity: 0.00 which is mostly neutral 
The first 5 sentences are: 
Sentence #0: nytimes.comPlease enable JS and disable any ad blocker 

**Sample Output 2:**  
Text from https://www.cbsnews.com/news/ketanji-brown-jackson-supreme-court-senate-confirmation-first-black-woman/ has results:  
Reability Dale-Chall score: 10.74   
Sentiment Polarity: 0.13 which is mostly neutral   
The first 5 sentences are:   
Sentence #0: Ketanji Brown Jackson confirmed to Supreme Court, making history as first Black female justiceHawaii WildfiresLGBTQ+ Veterans' LawsuitMega Millions WinnerNew COVID ShotsCBS News LiveManaging Your MoneyNewslettersNewsUSWorldPoliticsEntertainmentHealthWatchMoneyWatchCBS VillageTechnologyScienceCrimeSportsEssentialsAt least 6 killed on Hawaii's Maui island as fires force people to fleeNewly unveiled memo cited in Trump indictment detailed false electors schemeMan killed during FBI raid made threats against Biden, sources sayRobbie Robertson, guitarist for The Band, dies at age 80U.S. 
Sentence #1: nurse and daughter released more than a week after kidnap in HaitiTop doctor leaving Louisiana due to anti-LGBTQ legislationSpecial counsel obtained search warrant for Trump's Twitter accountFeinstein recovering after hospital visit for "minor fall" at California homeWhat to know about the new COVID shots this fallShowsCBS MorningsCBS Evening News60 MinutesCBS Saturday MorningFace the NationSunday Morning48 HoursCBS ReportsCBS News MorningsAmerica DecidesPrime Time with John DickersonCBS News WeekenderThe TakeoutThe UpliftEye On AmericaThe DishHere Comes the SunPerson To PersonLiveCBS News LiveCBS News BaltimoreCBS News Bay AreaCBS News BostonCBS News ChicagoCBS News ColoradoCBS News DetroitCBS News Los AngelesCBS News MiamiCBS News MinnesotaCBS News New YorkCBS News PhiladelphiaCBS News PittsburghCBS News SacramentoCBS News TexasCBS Sports HQMixibleLocalAtlantaBaltimoreBay AreaBostonChicagoColoradoDetroitLos AngelesMiamiMinnesotaNew YorkPhiladelphiaPittsburghSacramentoSeattleTampaTexasMoreLatestVideoPhotosPodcastsIn DepthLocalA Moment With…Innovators & DisruptorsLog InNewslettersMobileRSSCBS StoreParamount+Join Our Talent CommunityDavos 2023SearchSearchLoginWatch CBS NewsKetanji Brown Jackson confirmed to Supreme Court, making history as first Black female justiceget the free apphttps://www.cbsnews.com/news/ketanji-brown-jackson-supreme-court-senate-confirmation-first-black-woman/link copiedBy Melissa QuinnUpdated on:  April 7, 2022 / 7:13 PM/ CBS NewsFirst Black woman confirmed to Supreme CourtKetanji Brown Jackson becomes first Black woman confirmed to Supreme Court03:19Washington —The Senate voted to confirmJudge Ketanji Brown Jacksonto the Supreme Court on Thursday, cementing her place in history as the first Black woman to serve on the nation's highest court.Jackson's confirmation as the 116th justice in U.S. history received bipartisan backing, with a final vote of 53 to 47 in the upper chamber.Three Republicans, Sens. 
Sentence #2: Susan Collins of Maine, Lisa Murkowski of Alaska and Mitt Romney of Utah, joined all 50 Democrats in supporting President Biden's nominee. 
Sentence #3: Vice President Kamala Harris, the first woman and first woman of color to hold the role, presided over the Senate during the vote. 
Sentence #4: "On this vote, the yays are 53. 


**Sample Output 3:**  
Text from https://www.cbsnews.com/news/ketanji-brown-jackson-supreme-court-senate-confirmation-first-black-woman/ has results:   
Reability Dale-Chall score: 10.74   
Sentiment Polarity: 0.13 which is mostly neutral   
The first 5 sentences are:   
Sentence #0: Ketanji Brown Jackson confirmed to Supreme Court, making history as first Black female justiceHawaii WildfiresLGBTQ+ Veterans' LawsuitMega Millions WinnerNew COVID ShotsCBS News LiveManaging Your MoneyNewslettersNewsUSWorldPoliticsEntertainmentHealthWatchMoneyWatchCBS VillageTechnologyScienceCrimeSportsEssentialsAt least 6 killed on Hawaii's Maui island as fires force people to fleeNewly unveiled memo cited in Trump indictment detailed false electors schemeMan killed during FBI raid made threats against Biden, sources sayRobbie Robertson, guitarist for The Band, dies at age 80U.S. 
Sentence #1: nurse and daughter released more than a week after kidnap in HaitiTop doctor leaving Louisiana due to anti-LGBTQ legislationSpecial counsel obtained search warrant for Trump's Twitter accountFeinstein recovering after hospital visit for "minor fall" at California homeWhat to know about the new COVID shots this fallShowsCBS MorningsCBS Evening News60 MinutesCBS Saturday MorningFace the NationSunday Morning48 HoursCBS ReportsCBS News MorningsAmerica DecidesPrime Time with John DickersonCBS News WeekenderThe TakeoutThe UpliftEye On AmericaThe DishHere Comes the SunPerson To PersonLiveCBS News LiveCBS News BaltimoreCBS News Bay AreaCBS News BostonCBS News ChicagoCBS News ColoradoCBS News DetroitCBS News Los AngelesCBS News MiamiCBS News MinnesotaCBS News New YorkCBS News PhiladelphiaCBS News PittsburghCBS News SacramentoCBS News TexasCBS Sports HQMixibleLocalAtlantaBaltimoreBay AreaBostonChicagoColoradoDetroitLos AngelesMiamiMinnesotaNew YorkPhiladelphiaPittsburghSacramentoSeattleTampaTexasMoreLatestVideoPhotosPodcastsIn DepthLocalA Moment With…Innovators & DisruptorsLog InNewslettersMobileRSSCBS StoreParamount+Join Our Talent CommunityDavos 2023SearchSearchLoginWatch CBS NewsKetanji Brown Jackson confirmed to Supreme Court, making history as first Black female justiceget the free apphttps://www.cbsnews.com/news/ketanji-brown-jackson-supreme-court-senate-confirmation-first-black-woman/link copiedBy Melissa QuinnUpdated on:  April 7, 2022 / 7:13 PM/ CBS NewsFirst Black woman confirmed to Supreme CourtKetanji Brown Jackson becomes first Black woman confirmed to Supreme Court03:19Washington —The Senate voted to confirmJudge Ketanji Brown Jacksonto the Supreme Court on Thursday, cementing her place in history as the first Black woman to serve on the nation's highest court.Jackson's confirmation as the 116th justice in U.S. history received bipartisan backing, with a final vote of 53 to 47 in the upper chamber.Three Republicans, Sens. 
Sentence #2: Susan Collins of Maine, Lisa Murkowski of Alaska and Mitt Romney of Utah, joined all 50 Democrats in supporting President Biden's nominee. 
Sentence #3: Vice President Kamala Harris, the first woman and first woman of color to hold the role, presided over the Senate during the vote. 
Sentence #4: "On this vote, the yays are 53. 


## Let's start coding!

#### Using this Notebook 

Notice that some of the cells are not able to be edited – they are just meant to be run to test your code, and perform the auto grading. You should begin at the top and execute each code cell, adding your code when you see # YOUR CODE HERE. If you have written the functions correctly, you should see no errors.

When you execute the TEST CASES cell, what you want to see is no output. You may want to intentionally make the function give a wrong result, so that you can see what the error message looks like – and then fix your function. Once you execute the TEST CASES cell, and get no output, then your function is most likely correct!

```python
import requests, inspect
from bs4 import BeautifulSoup
from textblob import TextBlob
from textatistic import Textatistic
```

### Step 1: You are provided the function report_pol

As this function is provided in advance, it is readonly. You do not need to modify this cell, but you do need to run it.

```python
### UNGRADED FUNCTION

def report_pol(pol):
    ''' Report the polarity as positive, negative, or neutral '''
    if pol < -0.60:
        result = f'very negative'
    elif pol < -0.35:
        result = f'somewhat negative'
    elif pol < 0.35:
        result = f'mostly neutral'
    elif pol < 0.60:
        result = f'somewhat positive'
    else :
        result = f'positive'
    return result 
```

### Step 2: Write function extract_text(addr)

<!-- grade_id: cell-b8df4cc0186a657c, points: 0, grade: false, solution: true -->
```python
### GRADED FUNCTION:  Write a second function called "extract_text"

def extract_text(addr):
    ''' visit the website in the address, extract the text using Beautiful Soup,
        remove the html tags and return the text'''
    response = requests.get(addr)
    soup = BeautifulSoup(response.content, 'html.parser')
    text = soup.get_text(strip=True)
    return text
```

```python
# Try out your function
addr = "https://www.cbsnews.com/news/ketanji-brown-jackson-supreme-court-senate-confirmation-first-black-woman/"
text = extract_text(addr)
print(text)
```

### Step 3: Write function print_sentences(text, num_sentences)

<!-- grade_id: cell-246433127d09a89f, points: 0, grade: false, solution: true -->
```python
### GRADED FUNCTION:  Write a third function called "print_sentences"

def print_sentences(text, num_sentences):
    ''' print the first few sentences in the text, where the number of sentences
        is given by num_sentences '''
    blob = TextBlob(text)
    for i in range(num_sentences):
        print(blob.sentences[i])
    
```

```python
# Try out your function
addr = "https://www.cbsnews.com/news/ketanji-brown-jackson-supreme-court-senate-confirmation-first-black-woman/"
text = extract_text(addr)
print(print_sentences(text,2))
```

Output:
```
Ketanji Brown Jackson confirmed to Supreme Court, making history as first Black female justiceHunter Biden LawsuitNew Orleans WaterMaui Fire AftermathTrump InvestigationsCBS News LiveManaging Your MoneyNewslettersNewsUSWorldPoliticsEntertainmentHealthWatchMoneyWatchCBS VillageTechnologyScienceCrimeSportsEssentialsDonald Trump and his company "repeatedly" violated fraud law, N.Y. judge rulesHusband of NYC day care owner arrested in child's fentanyl deathHunter Biden sues Rudy Giuliani, attorney Robert Costello for "hacking" laptopAt least 20 dead in gas station blast as Nagorno-Karabakh residents fleeTarget says it's closing 9 stores because of surging retail theftsFTC and 17 states file sweeping antitrust suit against AmazonSupreme Court denies Alabama's bid to use GOP-drawn congressional map"Los Chapitos" Mexican cartel members sanctioned for fentanyl traffickingPolice investigating death of former NFL wide receiver Mike WilliamsShowsCBS MorningsCBS Evening News60 MinutesCBS Saturday 
```

### Step 4: Write function readability(text)

<!-- grade_id: cell-4488f95397e1c043, points: 0, grade: false, solution: true -->
```python
### GRADED FUNCTION:  Write a fifth function for option 4 called "generate_sentiment"

def readability(text):
    ''' return the readability score reported by Textastistic for Dale-Chall '''
    readability = Textatistic(text)
    result = readability.dalechall_score 
    return result
```

```python
# Try out your function
addr = "https://www.cbsnews.com/news/ketanji-brown-jackson-supreme-court-senate-confirmation-first-black-woman/"
text = extract_text(addr)
readability(text)
```

Output:
```
10.749461333121053
```

### Step 5: Write function sentiment(text)

<!-- grade_id: cell-f4c6e56f60f4084c, points: 0, grade: false, solution: true -->
```python
def sentiment(text):
    ''' return the sentiment score reported by TextBlob '''
    blob = TextBlob(text)
    return blob.sentiment.polarity
```

```python
# Try out your function
addr = "https://www.cbsnews.com/news/ketanji-brown-jackson-supreme-court-senate-confirmation-first-black-woman/"
text = extract_text(addr)
sentiment(text)
```

Output:
```
0.13183955400341954
```

### Step 6: Bring it all together in main()

```python
def main():
# set three websites to compare
    website1 = 'https://www.nytimes.com/live/2022/04/07/us/ketanji-brown-jackson-vote-scotus' 
    website2 = 'https://www.cbsnews.com/news/ketanji-brown-jackson-supreme-court-senate-confirmation-first-black-woman/'
    website3 = 'https://www.foxnews.com/politics/senate-confirmation-vote-ketanji-brown-jackson-supreme-court'

    # get data using requests, then process using BeautifulSoup
    text1 = extract_text(website1)
    text2 = extract_text(website2)
    text3 = extract_text(website3)

    #######
    # Load the text into the Textatistic library, and extract the Dale-Chall 
    # readability score.
    #######
    readability1 = readability(text1)
    readability2 = readability(text2)
    readability3 = readability(text3)

    #######
    # Load the text into the TextBlob library, and extract the sentiment polarity, 
    #######
    sentiment1 = sentiment(text1)
    sentiment2 = sentiment(text2)
    sentiment3 = sentiment(text3)

    #######
    # Report results
    #######
    print(f'Text from {website1} has results:')
    print(f'Reability Dale-Chall score: {readability1:0.2f}')
    print(f'Sentiment Polarity: {sentiment1:0.2f} which is {report_pol(sentiment1)}')
    print('The first 5 sentences are:')
    print_sentences(text1, 5)
    print('\n')
    print(f'Text from {website2} has results:')
    print(f'Reability Dale-Chall score: {readability2:0.2f}')
    print(f'Sentiment Polarity: {sentiment2:0.2f} which is {report_pol(sentiment2)}')
    print('The first 5 sentences are:')
    print_sentences(text2, 5)
    print('\n')
    print(f'Text from {website3} has results:')
    print(f'Reability Dale-Chall score: {readability3:0.2f}')
    print(f'Sentiment Polarity: {sentiment3:0.2f} which is {report_pol(sentiment3)}')
    print('The first 5 sentences are:')
    print_sentences(text3, 5)
    print('\n')
    
    # you might want to uncomment this and take a closer look at the text 
    #print('\n\n\nThe entire second article is:')
    #print(text2)
            
main()
```

Output:
```
Text from https://www.nytimes.com/live/2022/04/07/us/ketanji-brown-jackson-vote-scotus has results:
Cell In[63], line 53
     47     print('\n')
     49     # you might want to uncomment this and take a closer look at the text 
     50     #print('\n\n\nThe entire second article is:')
     51     #print(text2)
---> 53 main()

Cell In[63], line 31, in main()
     27 #######
     28 # Report results
     29 #######
     30 print(f'Text from {website1} has results:')
---> 31 print(f'Reability Dale-Chall score: {readability1:0.2f}')
     32 print(f'Sentiment Polarity: {sentiment1:0.2f} which is {report_pol(sentiment1)}')
     33 print('The first 5 sentences are:')

TypeError: unsupported format string passed to Textatistic.__format__
```

### Step 7: Testing Your Work

Your work will be autograded based on your function outputs for the test cases below. The next cell shows how the auto-grading tests are executed.  

Several possible cases are run through the function, and the results are tested against that actual correct answer. If you run the test cell and there is no output, then the tests all passed – Great!

<!-- grade_id: cell-19240d2478ff1f89, points: 3, grade: true, solution: false -->
```python
# GRADED TEST CASE 1: This cell should not deliver any errors when this cell is run
# Function 1: extract_text(addr)

# Test 1: requests, BeautifulSoup, html tag removal in syntax
req_bs4_found = False
tags_removed = False

list_time_lines = inspect.getsource(extract_text)
code_lines = list_time_lines.splitlines()
for line in code_lines:
    if "requests" in line or "BeautifulSoup" in line:
        req_bs4_found = True
    elif "get_text(strip=True)" in line:
        tags_removed  = True
    else:
        continue
        
# Test 2:
site_data_success = False

try:
    addr = "https://en.wikipedia.org/wiki/Ball_State_University"
    results = extract_text(addr)
    site_data_success = True
except Exception as e:
    print(e)
    
assert req_bs4_found == True, "Please verify that you've used requests and BeautifulSoup to acquire site data."
assert tags_removed == True, "Please verifiy that your function removes HTML tags with soup commands."
assert site_data_success == True, "Please verifiy that your function reaches a valid url and returns appropriate site data."
```

<!-- grade_id: cell-972625f7aab31e18, points: 4, grade: true, solution: false -->
```python
# GRADED TEST CASE 2: This cell should not deliver any errors when this cell is run

# Function 2 print_sentences(text, num_sentences):
sample_html = "Ketanji Brown Jackson confirmed to Supreme Court, making history as first Black female justiceJacksonville ShootingTropical Storm IdaliaTrump InvestigationsAI In ClassroomsCBS News LiveManaging Your MoneyNewslettersNewsUSWorldPoliticsEntertainmentHealthWatchMoneyWatchCBS VillageTechnologyScienceCrimeSportsEssentialsHurricane Idalia path, timeline: When and where storm will hit FloridaJudge sets March 2024 trial date in Trump's federal election caseHere's what experts say to know about new COVID variant BA.2.86Former Trump chief of staff Mark Meadows takes the stand in Georgia case'Armed and dangerous' person reported on UNC campusDeSantis booed at vigil for Jacksonville shooting victimsUgandan man faces possible death penalty under draconian anti-gay lawJoe the Plumber, who questioned Obama on his tax plans, dead at49Student loan repayments are set to resume. Here's what to know.ShowsCBS MorningsCBS Evening News60 MinutesCBS Saturday MorningFace the NationSunday Morning48 HoursCBS ReportsCBS News MorningsAmerica DecidesPrime Time with John DickersonCBS News WeekenderThe TakeoutThe UpliftEye On AmericaThe DishHere Comes the SunPerson To PersonLiveCBS News LiveCBS News BaltimoreCBS News Bay AreaCBS News BostonCBS News ChicagoCBS News ColoradoCBS News DetroitCBS News Los AngelesCBS News MiamiCBS News MinnesotaCBS News New YorkCBS News PhiladelphiaCBS News PittsburghCBS News SacramentoCBS News TexasCBS Sports HQMixibleLocalAtlantaBaltimoreBay AreaBostonChicagoColoradoDetroitLos AngelesMiamiMinnesotaNew YorkPhiladelphiaPittsburghSacramentoSeattleTampaTexasMoreLatestVideoPhotosPodcastsIn DepthLocalA Moment With…Innovators & DisruptorsLog InNewslettersMobileRSSCBS StoreParamount+Join Our Talent CommunityDavos 2023SearchSearchLoginWatch CBS NewsKetanji Brown Jackson confirmed to Supreme Court, making history as first Black female justiceget the free apphttps://www.cbsnews.com/news/ketanji-brown-jackson-supreme-court-senate-confirmation-first-black-woman/link copiedBy Melissa QuinnUpdated on:  April 7, 2022 / 7:13 PM/ CBS NewsFirst Black woman confirmed to Supreme CourtKetanji Brown Jackson becomes first Black woman confirmed to Supreme Court03:19Washington —The Senate voted to confirmJudge Ketanji Brown Jacksonto the Supreme Court on Thursday, cementing her place in history as the first Black woman to serve on the nation's highest court.Jackson's confirmation as the 116th justice in U.S. history received bipartisan backing, with a final vote of 53 to 47 in the upper chamber.Three Republicans, Sens. Susan Collins of Maine, Lisa Murkowski of Alaska and Mitt Romney of Utah, joined all 50 Democrats in supporting President Biden's nominee. Vice President Kamala Harris, the first woman and first woman of color to hold the role, presided over the Senate during the vote.'On this vote, the yays are 53. The nays are 47 and this nomination is confirmed,' Harris said to rousing applause from senators."

sentence_length = 2
sentence_data = print_sentences(sample_html,sentence_length)

# Test None is returned
assert sentence_data == None, "Please verify that your print_sentences function does not return any data, but prints a desired number of sentences."

# Test sentences, TextBlob, and print statement used:
textblob_check = False
print_check = False
sentences_check = False
loop_check = False

list_time_lines = inspect.getsource(print_sentences)
code_lines = list_time_lines.splitlines()

for line in code_lines:
    if "TextBlob" in line:
        textblob_check = True
    elif "print" in line:
        print_check = True
        
    if "sentence" in line:
        sentences_check = True
    if "for" in line and "in" in line:
        loop_check = True

assert textblob_check == True, "Please verify you've used TextBlob in your print_sentences function"
assert print_check == True,  "Please verify you've printed out sentences in your print_sentences function"
assert sentences_check == True, "Please verify you've isolated sentences in your print_sentences function"
assert loop_check == True, "Please verify you've looped through sentences as part of your print_sentences function"
```

Output:
```
Ketanji Brown Jackson confirmed to Supreme Court, making history as first Black female justiceJacksonville ShootingTropical Storm IdaliaTrump InvestigationsAI In ClassroomsCBS News LiveManaging Your MoneyNewslettersNewsUSWorldPoliticsEntertainmentHealthWatchMoneyWatchCBS VillageTechnologyScienceCrimeSportsEssentialsHurricane Idalia path, timeline: When and where storm will hit FloridaJudge sets March 2024 trial date in Trump's federal election caseHere's what experts say to know about new COVID variant BA.2.86Former Trump chief of staff Mark Meadows takes the stand in Georgia case'Armed and dangerous' person reported on UNC campusDeSantis booed at vigil for Jacksonville shooting victimsUgandan man faces possible death penalty under draconian anti-gay lawJoe the Plumber, who questioned Obama on his tax plans, dead at49Student loan repayments are set to resume.
Here's what to know.ShowsCBS MorningsCBS Evening News60 MinutesCBS Saturday MorningFace the NationSunday Morning48 HoursCBS Repo
```

<!-- grade_id: cell-3a25b6898afbf1a6, points: 2, grade: true, solution: false -->
```python
# GRADED TEST CASE 3: This cell should not deliver any errors when this cell is run

# Function readability(text)
sample_text = "Ketanji Brown Jackson confirmed to Supreme Court, making history as first Black female justiceJacksonville ShootingTropical Storm IdaliaTrump InvestigationsAI In ClassroomsCBS News LiveManaging Your MoneyNewslettersNewsUSWorldPoliticsEntertainmentHealthWatchMoneyWatchCBS VillageTechnologyScienceCrimeSportsEssentialsHurricane Idalia path, timeline: When and where storm will hit FloridaJudge sets March 2024 trial date in Trump's federal election caseHere's what experts say to know about new COVID variant BA.2.86Former Trump chief of staff Mark Meadows takes the stand in Georgia case'Armed and dangerous' person reported on UNC campusDeSantis booed at vigil for Jacksonville shooting victimsUgandan man faces possible death penalty under draconian anti-gay lawJoe the Plumber, who questioned Obama on his tax plans, dead at49Student loan repayments are set to resume. Here's what to know.ShowsCBS MorningsCBS Evening News60 MinutesCBS Saturday MorningFace the NationSunday Morning48 HoursCBS ReportsCBS News MorningsAmerica DecidesPrime Time with John DickersonCBS News WeekenderThe TakeoutThe UpliftEye On AmericaThe DishHere Comes the SunPerson To PersonLiveCBS News LiveCBS News BaltimoreCBS News Bay AreaCBS News BostonCBS News ChicagoCBS News ColoradoCBS News DetroitCBS News Los AngelesCBS News MiamiCBS News MinnesotaCBS News New YorkCBS News PhiladelphiaCBS News PittsburghCBS News SacramentoCBS News TexasCBS Sports HQMixibleLocalAtlantaBaltimoreBay AreaBostonChicagoColoradoDetroitLos AngelesMiamiMinnesotaNew YorkPhiladelphiaPittsburghSacramentoSeattleTampaTexasMoreLatestVideoPhotosPodcastsIn DepthLocalA Moment With…Innovators & DisruptorsLog InNewslettersMobileRSSCBS StoreParamount+Join Our Talent CommunityDavos 2023SearchSearchLoginWatch CBS NewsKetanji Brown Jackson confirmed to Supreme Court, making history as first Black female justiceget the free apphttps://www.cbsnews.com/news/ketanji-brown-jackson-supreme-court-senate-confirmation-first-black-woman/link copiedBy Melissa QuinnUpdated on:  April 7, 2022 / 7:13 PM/ CBS NewsFirst Black woman confirmed to Supreme CourtKetanji Brown Jackson becomes first Black woman confirmed to Supreme Court03:19Washington —The Senate voted to confirmJudge Ketanji Brown Jacksonto the Supreme Court on Thursday, cementing her place in history as the first Black woman to serve on the nation's highest court.Jackson's confirmation as the 116th justice in U.S. history received bipartisan backing, with a final vote of 53 to 47 in the upper chamber.Three Republicans, Sens. Susan Collins of Maine, Lisa Murkowski of Alaska and Mitt Romney of Utah, joined all 50 Democrats in supporting President Biden's nominee. Vice President Kamala Harris, the first woman and first woman of color to hold the role, presided over the Senate during the vote.'On this vote, the yays are 53. The nays are 47 and this nomination is confirmed,' Harris said to rousing applause from senators."
readability(sample_text)

# Verify dalechall_score used
dalechall_score_check = False

list_time_lines = inspect.getsource(readability)
code_lines = list_time_lines.splitlines()

for line in code_lines:
    if "dalechall_score" in line:
        dalechall_score_check = True

assert readability(sample_text) == 12.26130911567996, "Please verify you've implemented your readability function correctly"
assert dalechall_score_check == True, "Please verify you've implemented your function with the dalechall_score"
```

<!-- grade_id: cell-5ff291336958beff, points: 1, grade: true, solution: false -->
```python
# GRADED TEST CASE 4: This cell should not deliver any errors when this cell is run
sample_text = "Ketanji Brown Jackson confirmed to Supreme Court, making history as first Black female justiceJacksonville ShootingTropical Storm IdaliaTrump InvestigationsAI In ClassroomsCBS News LiveManaging Your MoneyNewslettersNewsUSWorldPoliticsEntertainmentHealthWatchMoneyWatchCBS VillageTechnologyScienceCrimeSportsEssentialsHurricane Idalia path, timeline: When and where storm will hit FloridaJudge sets March 2024 trial date in Trump's federal election caseHere's what experts say to know about new COVID variant BA.2.86Former Trump chief of staff Mark Meadows takes the stand in Georgia case'Armed and dangerous' person reported on UNC campusDeSantis booed at vigil for Jacksonville shooting victimsUgandan man faces possible death penalty under draconian anti-gay lawJoe the Plumber, who questioned Obama on his tax plans, dead at49Student loan repayments are set to resume. Here's what to know.ShowsCBS MorningsCBS Evening News60 MinutesCBS Saturday MorningFace the NationSunday Morning48 HoursCBS ReportsCBS News MorningsAmerica DecidesPrime Time with John DickersonCBS News WeekenderThe TakeoutThe UpliftEye On AmericaThe DishHere Comes the SunPerson To PersonLiveCBS News LiveCBS News BaltimoreCBS News Bay AreaCBS News BostonCBS News ChicagoCBS News ColoradoCBS News DetroitCBS News Los AngelesCBS News MiamiCBS News MinnesotaCBS News New YorkCBS News PhiladelphiaCBS News PittsburghCBS News SacramentoCBS News TexasCBS Sports HQMixibleLocalAtlantaBaltimoreBay AreaBostonChicagoColoradoDetroitLos AngelesMiamiMinnesotaNew YorkPhiladelphiaPittsburghSacramentoSeattleTampaTexasMoreLatestVideoPhotosPodcastsIn DepthLocalA Moment With…Innovators & DisruptorsLog InNewslettersMobileRSSCBS StoreParamount+Join Our Talent CommunityDavos 2023SearchSearchLoginWatch CBS NewsKetanji Brown Jackson confirmed to Supreme Court, making history as first Black female justiceget the free apphttps://www.cbsnews.com/news/ketanji-brown-jackson-supreme-court-senate-confirmation-first-black-woman/link copiedBy Melissa QuinnUpdated on:  April 7, 2022 / 7:13 PM/ CBS NewsFirst Black woman confirmed to Supreme CourtKetanji Brown Jackson becomes first Black woman confirmed to Supreme Court03:19Washington —The Senate voted to confirmJudge Ketanji Brown Jacksonto the Supreme Court on Thursday, cementing her place in history as the first Black woman to serve on the nation's highest court.Jackson's confirmation as the 116th justice in U.S. history received bipartisan backing, with a final vote of 53 to 47 in the upper chamber.Three Republicans, Sens. Susan Collins of Maine, Lisa Murkowski of Alaska and Mitt Romney of Utah, joined all 50 Democrats in supporting President Biden's nominee. Vice President Kamala Harris, the first woman and first woman of color to hold the role, presided over the Senate during the vote.'On this vote, the yays are 53. The nays are 47 and this nomination is confirmed,' Harris said to rousing applause from senators."

# Function 4 sentiment(text):
assert sentiment(sample_text) == 0.10331088664421999, "Please verify you've implemented your sentiment function accurately"
```

### Step 8: Submit your Work for Grading

To receive a final score for your work, please select the "Submit Assignment" button at the top of your lab.