# Programming Assignment 7a

## Assignment Instructions

Write two functions, described below.  Be sure to test your code against the test cases! Submit your code once your tests are passing.   Note: It is possible to earn partial credit for passing some of the tests, but not passing others.

* #### Function 1: list_times_100(orig_list) 
Write a function that is passed a list of numbers.  This function should use a for loop to multiply each list element by 100, creating a new list.  Return this new list. 

* #### Function 2: Function 1: list_times_100(orig_list) 
Now, write a second function that meets the same specifications as above, but instead uses a list comprehension to do the new list creation. 

### Optional Extension

For this assignment, you are provided with a main() function that will do some simple tests on these two functions.  If you’d like, modify main() so that you instead prompt the user to enter 7 numbers, and store them in a list.  Afterwards, call the two functions and print the results.  For additional challenge, you can ask the user how many numbers they want to enter, instead of assuming they always want to enter 7. 

## Let's start coding!

### Step 1: Write Function A

Notice that some of the cells are not able to be edited – they are just meant to be run to test your code, and perform the auto grading.  You should begin at the top and execute each code cell, adding your code when you see # YOUR CODE HERE.  If you have written the functions correctly, you should see no errors.

When you execute the TEST CASES cell, what you want to see is no output.  You may want to intentionally make the function give a wrong result, so that you can see what the error message looks like – and then fix your function.  Once you execute the TEST CASES cell, and get no output, then your function is most likely correct!


```python
## import all dependencies
import inspect 
```

<!-- grade_id: cell-44a0994fc0d12898, points: 0, grade: false, solution: true -->
```python
### GRADED FUNCTION A: Write a function "list_times_100"

def list_times_100(orig_list):
    ''' use a for loop to create a new list that is 100 times the old list '''
    total = []
    for each_number in orig_list:
        total.append(each_number * 100)
    return total

```

```python
## Try out your function with a simple list input
orig_list = [1,1,1,1]
list_times_100(orig_list)

print("Your list times 100 is: ", list_times_100(orig_list) )
```

Output:
```
Your list times 100 is:  [100, 100, 100, 100]
```

### Step 2: Write Function B

<!-- grade_id: cell-58341e91ad0b7cba, points: 0, grade: false, solution: true -->
```python
### GRADED FUNCTION B:  Write a second function called "list_comprehension_100"

def list_comprehension_100(starting_list):
    ''' use a list comprehension to create a new list that is 100 times the old list '''
    total = [(each_number * 100) for each_number in starting_list]
    return total
```

```python
## Try out your list comprehension function with a simple list input
starting_list = [2,2,2,2]
list_comprehension_100(starting_list)

print("Your new list times 100 is: ", list_comprehension_100(starting_list) )
```

Output:
```
Your new list times 100 is:  [100, 100, 100, 100]
```

```python
def main():
    print('Testing the two functions using two different lists:')
    
    # first list
    list1 = [10, 20, 30, 40]
    print(f'Starting with a list: {list1}')
    print(f'Using a for loop, we see: {list_times_100(list1)}')
    print(f'Using a list comprehension, we see: {list_comprehension_100(list1)}')
    print('They should be the same.')
    
    # second list
    list2 = [1, 2, 3, 4, 5, 6, 7, 8, 9, 10, 11]
    print(f'Starting with a list: {list2}')
    print(f'Using a for loop, we see: {list_times_100(list2)}')
    print(f'Using a list comprehension, we see: {list_comprehension_100(list1)}')
    print('They should be the same.')
    
main()
```

Output:
```
Testing the two functions using two different lists:
Starting with a list: [10, 20, 30, 40]
Using a for loop, we see: [1000, 2000, 3000, 4000]
Using a list comprehension, we see: [100, 100, 100, 100]
They should be the same.
Starting with a list: [1, 2, 3, 4, 5, 6, 7, 8, 9, 10, 11]
Using a for loop, we see: [100, 200, 300, 400, 500, 600, 700, 800, 900, 1000, 1100]
Using a list comprehension, we see: [100, 100, 100, 100]
They should be the same.
```

### Step 3: Testing Your Work

Below is a program that will use your functions to generate two new lists. Your work will be autograded based on your function outputs for the test cases below.

<!-- grade_id: cell-19240d2478ff1f89, points: 3, grade: true, solution: false -->
```python
# GRADED TEST CASES 1: This cell should not deliver any errors when this cell is run
print('Testing the two functions using two different lists:')

# first test list
list1 = [10, 20, 30, 40]
assert list_times_100(list1) == [1000, 2000, 3000, 4000], "Not quite. Verify that your for loop has the right multiplier. Your test output list is %s " % list_times_100(list1)
assert list_comprehension_100(list1) == [1000, 2000, 3000, 4000], "Not quite. Please verify your list comprehension logic. Your test output list is %s " % list_comprehension_100(list1)
assert list_times_100(list1) == list_comprehension_100(list1), "Not quite. %s should equal %s " % list_times_100(list1) % list_comprehension_100(list1)
```

Output:
```
Testing the two functions using two different lists:
```

<!-- grade_id: cell-972625f7aab31e18, points: 1, grade: true, solution: false -->
```python
# GRADED TEST CASES 2: This cell should not deliver any errors when this cell is run

# Check - for loop implementation
for_loop_found = False

list_time_lines = inspect.getsource(list_times_100)
code_lines = list_time_lines.splitlines()
for line in code_lines:
    if "for" in line and "in" in line and ":" in line:
        for_loop_found = True

assert for_loop_found == True, "Please verify you've implemented a for loop in Function A."
```

<!-- grade_id: cell-3a25b6898afbf1a6, points: 1, grade: true, solution: false -->
```python
# GRADED TEST CASES 3: This cell should not deliver any errors when this cell is run

## Check - list comprehension implementation
list_comprehension_found = False

list_comp_lines = inspect.getsource(list_comprehension_100)
comp_code_lines = list_comp_lines.splitlines()
for line in comp_code_lines :
    if "[" in line and "for" in line and "in starting_list" in line and "]" in line:
        list_comprehension_found = True

assert list_comprehension_found == True, "Please verify you've implemented a list comprehension in Function B."
```

<!-- grade_id: cell-2d949afda9933aad, points: 3, grade: true, solution: false -->
```python
# GRADED TEST CASE 4: This cell should not deliver any errors when this cell is run

# second test list
list2 = [1, 2, 3, 4, 5, 6, 7, 8, 9, 10, 11]
assert list_times_100(list2) == [100, 200, 300, 400, 500, 600, 700, 800, 900, 1000, 1100], "Not quite. Verify that your for loop has the right multiplier. Your test output list is %s " % list_times_100(list2)
assert list_comprehension_100(list2) == [100, 200, 300, 400, 500, 600, 700, 800, 900, 1000, 1100], "Not quite. Please verify your list comprehension logic. Your test output list is %s " % list_comprehension_100(list2)
assert list_times_100(list2) == list_comprehension_100(list2), "Not quite. %s should equal %s " % list_times_100(list2) % list_comprehension_100(list2)
```

### Step 4: Submit your Work for Grading

Great job on your work with for loops and list comprehension functions this week!

To receive a final score for your work, please select the "Submit Assignment" button at the top of your lab.