# Programming Assignment 4

## Assignment Instructions

Write two functions, described below.  Be sure to test your code against the test cases!  Submit your code once your tests are passing.  (It is possible to earn partial credit for passing some of the tests, but not passing others.) 

### __Function 1__:   

Define a function that will take no arguments, but will return an integer that represents the option selected (option 1, 2, 3, or 4).  The "inside" of this function should:  

Print a list of menu options that the user can select from.  Use these 4 options: 

1. See an uplifting message 

2. See a programming tip 

3. Watch me count 

4. Quit this program. 

Ask the user to enter a number to indicate their selection  

Use an if statement to see if they typed a 1.  If so, return the integer 1.    

*Hint:  you can compare strings, like this:*  

```python
if user_input == '1':  
    # they entered a 1, great!
```

In this way, you avoid the program crashing if they type "asdf" and then you try to convert that to an integer.   
Similarly, use elif statements to return 2, 3, and 4.  
Include an else statement that tells them that they entered an invalid number.  You can return a 0 in that case.  

### __Function 2__: 

Write a function with one argument that is an integer.  That function should use a for loop to print out the numbers from 1 to that number (inclusive).

### __Main Program__: *Putting it all together*

Now pull these two functions together into a whole program.  Your code should use your first function to present the user with a menu of options.  
If your first function returns a 1, your program will print an uplifting message (choose a quote that you find encouraging).  
If they instead pressed a 2, your program will print a programming tip (choose something that you have learned in this class).  
If they pressed 3, your program will then ask them for a number, then count up to that number, starting at 1 (using the function that you wrote).  
If they press 4, the program will exit.   

In this case, you do not need to continue looping – the user would need to run the program another time in order to choose a different option. 

When you run your program, it should produce output similar to what is shown below.  It does not need to be identical, but should be functionally the same. 

## Let's start coding!

### Using this Notebook

Notice that some of the cells are not able to be edited – they are just meant to be run to test your code, and perform the auto grading. You should begin at the top and execute each code cell, adding your code when you see # YOUR CODE HERE. If you have written the functions correctly, you should see no errors.

When you execute the TEST CASES cell, what you want to see is no output. You may want to intentionally make the function give a wrong result, so that you can see what the error message looks like – and then fix your function. Once you execute the TEST CASES cell, and get no output, then your function is most likely correct!

```python
# Install Necessary Dependencies
import nose, inspect
from nose.tools import assert_equal
```

### Function 1: menu

<!-- grade_id: cell-fb4aeb0d1c5c8ec4, points: 0, grade: false, solution: true -->
```python
def menu():
    ''' print a list of menu options, and return an integer representing the user's choice '''
    print('Choose between the options bellow:')
    print('1 - See an uplifting message.')
    print('2 - See a programming tip.') 
    print('3 - Watch me count.')
    print('4 - Quit this program.')
    choice = input('Write the number of your choice.')
    if choice == '1':
        return 1
    elif choice == '2':
        return 2
    elif choice == '3':
        return 3
    elif choice == '4':
        return 4
    else:
        print('You entered a invalid number.')
        return 0
    
    
menu()
```

Output:
```
Choose between the options bellow:
1 - See an uplifting message.
2 - See a programming tip.
3 - Watch me count.
4 - Quit this program.
Write the number of your choice. 2
2
```

```python
# MANUAL TESTING OF MENU
# This cell will ask for input 5 different times.  
# Enter each value (1, 2, 3, 4, 0 or any invalid value) one at a time.
assert menu() == 1, "Check your menu function with choice 1."
assert menu() == 2, "Check your menu function with choice 2."
assert menu() == 3, "Check your menu function with choice 3."
assert menu() == 4, "Check your menu function with choice 4."
assert menu() == 0, "Check your menu function with invalid choice."
```

Output:
```
Choose between the options bellow:
1 - See an uplifting message.
2 - See a programming tip.
3 - Watch me count.
4 - Quit this program.
Write the number of your choice. 1
Choose between the options bellow:
1 - See an uplifting message.
2 - See a programming tip.
3 - Watch me count.
4 - Quit this program.
Write the number of your choice. 2
Choose between the options bellow:
1 - See an uplifting message.
2 - See a programming tip.
3 - Watch me count.
4 - Quit this program.
Write the number of your choice. 3
Choose between the options bellow:
1 - See an uplifting message.
2 - See a programming tip.
3 - Watch me count.
4 - Quit this program.
Write the number of your choice. 4
Choose between the options bellow:
1 - See an uplifting message.
2 - See a programming tip.
3 - Watch me count.
4 - Quit this program.
Write the number of your choice. 5
You entered a invalid number.
```

<!-- grade_id: cell-b72fe6462a70755e, points: 5, grade: true, solution: false -->
```python
# GRADED TEST CASE 1: This cell should not deliver any errors when this cell is run
# Check - for loop implementation
if_statement_found = False
option_input_found = False
menu_option_found = False
return_val_found = False

menu_lines = inspect.getsource(menu)
code_lines = menu_lines.splitlines()
for line in code_lines:
    if "if" in line:
        if_statement_found = True
    if "input" in line:
        option_input_found = True
    if "see an uplifting message" in line.lower():
        menu_option_found = True
    if "return 4" in line:
        return_val_found = True
        
assert if_statement_found == True, "Please verify you've implemented if logic in your function."
assert option_input_found == True, "Please verify if you've asked the user for their input."
assert menu_option_found == True, "Please verify if you've implemented all menu options."
assert return_val_found == True, "Please verify if you've returned all potential integer selection options from 1-4."
```

### Function 2: count_up

<!-- grade_id: cell-95f796f66fb927cc, points: 0, grade: false, solution: true -->
```python
def count_up(num):
    ''' count from 1 to that number, inclusive '''
    for i in range(1, num + 1):
        print(i)
count_up(6)
```

Output:
```
1
2
3
4
5
6
```

```python
# MANUAL TESTING OF count_up
# This cell will ask for input 5 different times.  
# Enter each input value (1, 2, 3, 4, 5) one at a time.
assert count_up(1) == None, "Check your count_up function with number 1."
assert count_up(2) == None, "Check your count_up function with number 2."
assert count_up(3) == None, "Check your count_up function with number 3."
assert count_up(4) == None, "Check your count_up function with number 4."
assert count_up(5) == None, "Check your count_up function with number 5."
```

Output:
```
1
1
2
1
2
3
1
2
3
4
1
2
3
4
5
```

<!-- grade_id: cell-a296d0b0289e0378, points: 5, grade: true, solution: false -->
```python
# GRADED TEST CASE 2: This cell should not deliver any errors when this cell is run

# Check - for loop implementation
for_loop_found = False

count_up_lines = inspect.getsource(count_up)
code_lines = count_up_lines.splitlines()
for line in code_lines:
    if "for" in line and "in" in line and ":" in line:
        for_loop_found = True

assert for_loop_found == True, "Please verify you've implemented a for loop in your count_up function"
assert count_up(5) == None, "Check your count_up function with number 5."
```

Output:
```
1
2
3
4
5
```

### Main program: *Putting it all together*

<!-- grade_id: cell-7529a8dd611d2644, points: 0, grade: false, solution: true -->
```python
# Now pull these two functions together into a whole program.  
# Your code should use your first function to present the user with a menu of options.  
def main():
    while True:
        request = menu()
        if request == 1:
            print('Believe you can do and you are halfway there.')
            break
        elif request == 2:   
            print("Don't forget to use indentation in your function.")
            break
        elif request == 3:
            number = int(input('Enter a number: '))
            count_up(number)
            break
        elif request == 4:
            print('Exiting the program.')
            break
        else:
            continue                                                                                                          

# do not remove as this is calling the main function
main()
```

Output:
```
Choose between the options bellow:
1 - See an uplifting message.
2 - See a programming tip.
3 - Watch me count.
4 - Quit this program.
Write the number of your choice. 2
Don't forget to use indentation in your function.
```

```python

```

```python

```