# Programming Assignment 14

## Assignment Instructions

In this assignment, you will use the requests package and your understanding of JSON objects to build a very basic web browser in just a few lines of code.  We’ll limit it to connection only to Wikipedia – and only display limited information.  The user will be able to type in the exact Wikipedia article (for example, “Ball_State_University”).  Your program will connect to Wikipedia API, retrieve the JSON object, split out the information of interest, and print it to the screen. 

##### You should carefully study the following the Live Coding lectures: 

* “Using python to connect to an API: the requests package” shows how to use the requests package to connect to Wikipedia and retrieve a JSON object for an article. 

* “Exploring JSON structure” shows how to access complex internal JSON structures to extract the desired information, using an example JSON object. 

* You’ll need to use the techniques from the second Live Coding lecture to extract the information that you need from the Wikipedia JSON that is retrieved using the techniques for the first lecture. 

The function definitions for this lab are relatively general – and structured more like something you’d want to write for yourself, instead of an assignment.  However, we still need to make it gradable, so that you can have the instant feedback that is so important for learning – so we’ll define the function names, arguments, and return values for that purpose. 

- #### Function 1: get_Wikipedia_json_for_title(article_title) 
Write a function that will connect to Wikipedia and request an article of a specific title (passed to the function as a string). The function will return the JSON object received. 



- #### Function 2: json_title_and_toc(json_obj) 
Write a function that will accept a JSON object, and return a string that contains relevant information about the article sections.  The returned information must match the format shown below.  Each section will be returned as a line in the string, formatted exactly in this manner: 

    - *Section#: Title of section*

    - In the returned string, the # sign indicates the section number as stored in the JSON object, with no space between the word “Section” and the number. 

        - Is this totally arbitrary that I specify it this way?  Yes!  Does it help illustrate the point that when we work with data and write programs, we often must get something exactly correct in order for it to work?  Absolutely.  Very frequently, we encounter weird formatting requirements, and one character of a difference can cause the whole program to fail.  (In this case, it will not pass the grading requirements!) 

 

#### Main program 

As usual, main() is provided.  In this case, though, there is an annoying issue that should probably be fixed (but is optional!).  If you are up for a challenge, make it so that your program will NOT crash if you enter a title for an article that does not exist in Wikipedia (such as the article “asoidfjklwqer;ewqm”). 
 

#### Sample Outputs


**Sample Output 1:**
Please enter a wikipedia article title, with _ for spaces enter "quit" to exit:  
**User Input** The_Wizard_of_Oz   

The Wizard of Oz   
Section1: Adaptations of the novel   
Section1.1: Film   
Section1.2: Television and video   
Section1.3: Musicals and concerts   
Section1.4: Gaming   
Section2: People nicknamed "Wizard of Oz"   
Section3: Other uses   
Section4: See also   
 

**Sample Output 2:**  
Please enter a wikipedia article title, with _ for spaces enter "quit" to exit:  
**User Input** Zappa

Zappa

**Sample Output 3:**  
Please enter a wikipedia article title, with _ for spaces enter "quit" to exit:  
**User Input**  Frank_Zappa  

Frank Zappa 
Section1: 1940s–1960s: early life and career   
Section1.1: Childhood   
Section1.2: First musical interests   
Section1.3: Studio Z   
Section2: Late 1960s: the Mothers of Invention 
Section2.1: Formation    
Section2.2: Debut album: <i>Freak Out!</i>    
Section2.3: New York period (1966–1968)    
Section2.4: Disbandment    
Section3: 1970s    
Section3.1: Rebirth of the Mothers and filmmaking    
Section3.2: Accident, attack, and aftermath    
Section3.3: Solo albums: Waka/Jawaka and The Grand Wazoo    
Section3.4: Top 10 album: <i>Apostrophe (<span>&#39;</span>)</i>    
Section3.5: Business breakups and touring    
Section3.6: Zappa Records label    
Section4: 1980s–1990s    
Section4.1: "Valley Girl" and classical performances    
Section4.2: Synclavier works    
Section4.3: Merchandising    
Section4.4: Digital medium and last tour    
Section4.5: Health deterioration    
Section5: Death    
Section6: Musical style and development    
Section6.1: Genres    
Section6.2: Influences    
Section6.3: Project/Object    
Section6.4: Techniques    
Section6.4.1: Guitar playing    
Section6.4.2: Tape manipulation    
Section7: Personal life    
Section8: Beliefs and politics    
Section8.1: Drugs    
Section8.2: Government and religion    
Section8.3: Anti-censorship    
Section9: Legacy    
Section9.1: Acclaim and honors    
Section9.1.1: Grammy Awards    
Section9.2: Artists influenced by Zappa    
Section9.3: References in arts and sciences    
Section9.4: <i>Zappa</i> documentary    
Section10: Discography    
Section11: Tour    
Section12: See also    
Section13: Notes    
Section14: References    
Section15: Bibliography    
Section16: External links    

**Sample Output 4:**
Please enter a wikipedia article title, with _ for spaces enter "quit" to exit:  
**User Input** Python_(programming_language)  

Python (programming language)   
Section1: History   
Section2: Design philosophy and features   
Section3: Syntax and semantics   
Section3.1: Indentation   
Section3.2: Statements and control flow   
Section3.3: Expressions   
Section3.4: Methods   
Section3.5: Typing   
Section3.6: Arithmetic operations   
Section4: Programming examples   
Section5: Libraries   
Section6: Development environments   
Section7: Implementations   
Section7.1: Reference implementation   
Section7.2: Other implementations   
Section7.3: Unsupported implementations   
Section7.4: Cross-compilers to other languages   
Section7.5: Performance   
Section8: Development   
Section9: API documentation generators   
Section10: Naming   
Section11: Popularity   
Section12: Uses   
Section13: Languages influenced by Python   
Section14: See also   
Section15: References   
Section15.1: Sources   
Section16: Further reading   
Section17: External links   

**Sample Output 5:**
Please enter a wikipedia article title, with _ for spaces enter "quit" to exit:  
**User Input** Ball_State_University   

Ball State University    
Section1: History    
Section1.1: Predecessor schools    
Section1.2: The Ball brothers and Ball State Teachers College era (1917–1960)    
Section1.3: Independent institution (1961-present)    
Section2: Campus    
Section2.1: Main campus    
Section2.2: Campus transportation    
Section2.3: Athletic facilities    
Section2.4: Architecture       
Section2.5: Sustainability    
Section2.5.1: Geothermal system    
Section2.6: Satellite facilities    
Section3: Academics    
Section3.1: Student body    
Section3.2: Tuition    
Section3.3: Colleges    
Section3.4: Library system    
Section3.5: Rankings    
Section4: Student life    
Section4.1: Housing    
Section4.2: Student organizations and activities    
Section4.3: Media    
Section4.3.1: <i>Ball State Daily News</i>    
Section4.3.2: <i>Ball Bearings</i> magazine    
Section4.3.3: Byte    
Section4.3.4: Cardinal Metrics    
Section4.3.5: Cardinal WX    
Section4.3.6: NewsLink Indiana    
Section4.3.7: Ball State Sports Link    
Section4.3.8: WCRD    
Section4.3.9: The McKinley Avenue Agency    
Section4.3.10: Ball State Public Media    
Section5: Athletics    
Section6: Traditions    
Section6.1: <i>Beneficence</i>    
Section6.2: <i>Frog Baby</i>    
Section6.3: <i>The Naked Lady</i>    
Section6.4: Homecoming    
Section6.5: Other traditions    
Section7: Notable alumni    
Section8: See also    
Section9: Notes    
Section10: References    
Section11: Further reading    
Section12: External links   

Coursera**Sample Output 6:**  
Please enter a wikipedia article title, with _ for spaces enter "quit" to exit:  
**User Input**  Coursera  

Coursera   
Section1: History   
Section1.1: Finances   
Section1.2: Funding   
Section1.3: Business model   
Section2: Strategic partners   
Section3: Product and services   
Section3.1: Courses   
Section3.2: Degrees   
Section3.3: Professional certificate   
Section4: Impact   
Section4.1: Outreach and sponsorship programs   
Section5: References   
Section6: External links   

**Sample Output 7:**  
Please enter a wikipedia article title, with _ for spaces enter "quit" to exit:  
**User Input**  quit 

## Let's start coding!

#### Using this Notebook 

Notice that some of the cells are not able to be edited – they are just meant to be run to test your code, and perform the auto grading. You should begin at the top and execute each code cell, adding your code when you see # YOUR CODE HERE. If you have written the functions correctly, you should see no errors.

When you execute the TEST CASES cell, what you want to see is no output. You may want to intentionally make the function give a wrong result, so that you can see what the error message looks like – and then fix your function. Once you execute the TEST CASES cell, and get no output, then your function is most likely correct!

```python
## import all dependencies
import requests # for http connections
import json, inspect
```

### Step 1: Write function get_Wikipedia_json_for_title()

<!-- grade_id: cell-44a0994fc0d12898, points: 0, grade: false, solution: true -->
```python
### GRADED FUNCTION: Write the function "get_Wikipedia_json_for_title()"

def get_Wikipedia_json_for_title(article_title):
    ''' Write a function that will connect to Wikipedia and get an article 
    of a specific title, returning the json object '''
    url = 'https://en.wikipedia.org/w/api.php?action=parse&page=' + article_title +'&format=json'
    response = requests.get(url)
    json = response.json()      
    return json
```

```python
## Try out your function with a sample article title input
article_title = "The_Wizard_of_Oz"
wiki_data = get_Wikipedia_json_for_title(article_title)
print(wiki_data)
```

```python

```

### Step 2: Write function json_title_and_toc()

<!-- grade_id: cell-58341e91ad0b7cba, points: 0, grade: false, solution: true -->
```python
### GRADED FUNCTION:  Write a second function called "json_title_and_toc"

def json_title_and_toc(json_obj):
    ''' Write a function that will accept a json object, and 
    return a string that contains relevant information about the article sections '''
    json = json_obj['parse']['sections']
    string = str(json_obj['parse']['title'])
    for i in range (0, len(json)):
        string += f"\nSection{json[i]['number']}: {json[i]['line']}"
    return string
```

```python
## Try out your function using the short json_object returned from get_Wikipedia_json_for_title()
article_title = "Coursera"
wiki_data = get_Wikipedia_json_for_title(article_title)
summary = json_title_and_toc(wiki_data)

print(summary)
```

Output:
```
Coursera
Section1: History
Section1.1: Finances
Section1.2: Funding
Section1.3: Business model
Section2: Strategic partners
Section3: Product and services
Section3.1: Courses
Section3.2: Degrees
Section3.3: Professional certificate
Section4: Impact
Section4.1: Outreach and sponsorship programs
Section5: References
Section6: External links
```

```python
def main():
    while True:
        print('\nPlease enter a wikipedia article title, with _ for spaces')
        title = input('enter "quit" to exit: ')
        if title == 'quit':
            break
        json_response = get_Wikipedia_json_for_title(title)
        article_summary = json_title_and_toc(json_response)
        print(article_summary)
    
main()
```

### Step 3: Testing Your Work

Your work will be autograded based on your function outputs for the test cases below. The next cell shows how the auto-grading tests are executed.  

Several possible cases are run through the function, and the results are tested against that actual correct answer. If you run the test cell and there is no output, then the tests all passed – Great!

<!-- grade_id: cell-19240d2478ff1f89, points: 2, grade: true, solution: false -->
```python
# GRADED TEST CASES 1: This cell should not deliver any errors when this cell is run

# Function 1: Errors if parameter provided is not a valid string
article_title = 3
non_string_error = False

try:
    response = get_Wikipedia_json_for_title(article_title)
    non_string_error = False
except:
    non_string_error = True

assert non_string_error == True, "Please verify your function only takes string inputs"
```

<!-- grade_id: cell-972625f7aab31e18, points: 6, grade: true, solution: false -->
```python
# GRADED TEST CASES 2: This cell should not deliver any errors when this cell is run
# Function 1: Checks for a valid json response

json_url_syntax = False
returns_data_response = False
json_check = False
response_handled = False

#### Function returns data without error
try:
    response = get_Wikipedia_json_for_title("The_Wizard_of_Oz")
    returns_data_response = True
except:
    returns_data_response = False


## Checks that the url response is not the requests type
response = get_Wikipedia_json_for_title("Frank Zappa")

if type(response) is requests.models.Response:
    response_handled = False
else:
    response_handled = True

    
## Checks for json in dict from response.json()
try:
    response = get_Wikipedia_json_for_title("Python_(programming_language)")
    if type(response) is dict and "parse" in response:
        json_check = True
    data = json.dumps(response, indent=4, sort_keys=True) # check dict to json conversion with pretty print
    newdata = json.loads(data) # check json to python conversion
except:
    json_check = False
    
assert returns_data_response == True, "Please verify your function returns data correctly without erroring."
assert json_check == True, "Please verify data is returned in a json parsable format"
assert response_handled == True, "Please verify that you've returned the correct data type from your URL response"
```

<!-- grade_id: cell-3a25b6898afbf1a6, points: 2, grade: true, solution: false -->
```python
# GRADED TEST CASES 3: This cell should not deliver any errors when this cell is run
# Function 2: Checks for specific section formatting

def find_sections(article_title):
    wiki_data = get_Wikipedia_json_for_title(article_title)
    summary = json_title_and_toc(wiki_data)
    lines = summary.splitlines()
    count = 0
    sections = []

    for section in lines:
        if count != 0:
            if "Section" in section and ":" in section:
                sections.append(section)
        count += 1
    return sections

assert find_sections("The_Wizard_of_Oz") == ['Section1: Adaptations of the novel', 'Section1.1: Film', 'Section1.2: Television and video', 'Section1.3: Musicals and concerts', 'Section1.4: Gaming', 'Section2: People nicknamed "Wizard of Oz"', 'Section3: Other uses', 'Section4: See also'], "Please verify that you've returned each section in the following format 'Section#: Title of section'"
```

### Step 4: Submit your Work for Grading

To receive a final score for your work, please select the "Submit Assignment" button at the top of your lab.