# __Programming Assignment 6b__

## Assignment Instructions

#### The purpose of this assignment is to:

* Practice manipulating lists 

* Practice using for loops with lists 

* Practice sorting lists 

Follow these instructions to create a python program that will prompt the user to enter 5 words and add them to a list, but not allow duplicate words to be added. For this activity, you'll write two functions as described below. 

Be sure to test your code against the test cases! Submit your code once your tests are passing. (It is possible to earn partial credit for passing some of the tests, but not passing others.)

### __Function 1 :__ 

Write a function that accepts two arguments: a list of words, and a new word to add to the list, as long as it is not a duplicate.  

  1. If the word is already in the list, the function returns `True` (and does not add the word to the list a second time).
  2. If the word is not yet in the list, then the function adds it to the end of the list, and returns `False`. 


### __Function 2 :__


Write a function that accepts a list of words and returns a copy of the list where each word has been converted to Title Case, and then sorted.

##### Function 2 : Hints

- You may want to use a for loop to iterate through each word, processing it separately, and then adding it back to a new list.
- You can use the string function .title() to convert the word into title case.
- Remember to review the difference between .sort() and sorted() 
    - .sort() sorts a list and modifies the original list, returning the value None 
    - sorted() takes a list as an argument, sorts it, and returns a copy of the original list 

### __Main Program :__ *Putting it all together*

Finally, you'll bring both the functions together in a program that prompt the user to enter 5 words. If the user enters a duplicate, your program should tell them they have already entered that word. Then, print the original list. Then, sort the list and print it out a second time. 

#### Sample Output:

    Please enter a word:  Learning 
    Please enter a word:  python 
    Please enter a word:  is 
    Please enter a word:  very 
    Please enter a word:  very 

    You already entered that word. 

    Please enter a word:  very 

    You already entered that word. 

    Please enter a word:  very 

    You already entered that word. 

    Please enter a word:  fun! 

    You entered these words:  ['Learning', 'python', 'is', 'very', 'fun!'] 
    The sorted list in title case is:  ['Fun!', 'Is', 'Learning', 'Python', 'Very'] 

## Lets Start Coding! 

### Function 1 : 

<!-- grade_id: cell-0d76c9fef97c049a, points: 0, grade: false, solution: true -->
```python
def check_and_store(word_list, word):
    if word in word_list:
        return True
    else:
        word_list.append(word)
        return False
    
    

```

```python
# Ungraded: Try it out - test your function

word_list = ['ball']  # to test it, add a list of several words here, but include the word “ball” as one list element
word1 = "ball" # we want to test adding a duplicate word, so this word be the same as one in the list
word2 = "unique" # we want to test adding a new word, so this word should not be in the list

assert check_and_store(word_list, word1) == True , " Please check your function . A duplicate word should return True"
assert check_and_store(word_list, word2) == False , " Please check your function . A new word should be appended to the list and return false"
```

<!-- grade_id: cell-5e241d7cbd081526, points: 1, grade: true, solution: false -->
```python
#GRADED TEST CASES : Test your solution with the tests added in this cell

assert check_and_store(["Monday","Tuesday","Wednesday"], "Tuesday") == True, " Check your function . A duplicate word should return True"
assert check_and_store(["Monday","Tuesday","Wednesday"], "Saturday") == False, "Check your function . A new word appended to the list should return False"
```

### Function 2: 

<!-- grade_id: cell-4b5d3ddc99b3fcbc, points: 0, grade: false, solution: true -->
```python
def case_sorting(list_of_words):
    list_of_words = [word.capitalize() for word in list_of_words]
    list_of_words.sort()
    return list_of_words
    
case_sorting(['halef', 'renata', 'julia'])
```

Output:
```
['Halef', 'Julia', 'Renata']
```

```python
# Ungraded: Try it out - test your function

list_of_words1 = ['pear','orange','plum' ] # add a list of test words here
list_of_words2 = ['potato','radish','corn' ] # add a list of test words here

assert case_sorting(list_of_words1) == ['Orange','Pear','Plum'] , " Please check your function's sorting functionality"
assert case_sorting(list_of_words2) == ['Corn','Potato','Radish'] , " Please check your function's sorting functionality"
```

<!-- grade_id: cell-006de32726037f4c, points: 1, grade: true, solution: false -->
```python
# GRADED TEST CASE : Test your solution with the tests added in this cell
assert case_sorting(['fiction','autobiography','dictionary','motivational']) == ['Autobiography', 'Dictionary', 'Fiction', 'Motivational'], "There seems to be some error. Check your sorting function ! " 
assert case_sorting(['Mountains','rivers','ponds','lakes']) == ['Lakes', 'Mountains', 'Ponds', 'Rivers'], "There seems to be some error. Check your sorting function ! " 
```

### Main Program: *Putting it all together* 

Now lets bring it all together into one program through a main function. This part of the assignment is not graded but is highly recommended.  Examine the code for the main program below. Enter the code needed to make the program function as described in the comments # ADD CODE HERE.
 

<!-- grade_id: cell-7a89f05d1331b4a5, points: 0, grade: false, solution: true -->
```python
def check_and_store(word_list, word):
    if word in word_list:
        print('You already entered that word.')
        return True
    else:
        word_list.append(word)
        return False
    
    
def case_sorting(list_of_words):
    list_of_words = [word.capitalize() for word in list_of_words]
    list_of_words.sort()
    return list_of_words


def main():
    # loop to allow the user to enter 5 words
    list_of_words = []
    NUM_WORDS = 5
    while len(list_of_words) < NUM_WORDS:
        user_word = input('Please enter a word: ')
        # If the user enters a duplicate, tell them they have already entered that word
        # otherwise, store the new word in the list_of_words
        check_and_store(list_of_words, user_word)

    #print out the list 
    print('You entered these words: ', list_of_words)
    # convert the word list to title case and sort it
    print('The sorted list in title case is: ', case_sorting(list_of_words))

# Do not remove this line as it is calling the main function
main()

```

Output:
```
Please enter a word:  halef
Please enter a word:  renata
Please enter a word:  julia
Please enter a word:  debora
Please enter a word:  mae
You entered these words:  ['halef', 'renata', 'julia', 'debora', 'mae']
The sorted list in title case is:  ['Debora', 'Halef', 'Julia', 'Mae', 'Renata']
```

## Final Step : Submit your work for Grading

Great job on your work this week! To receive a final score for your work, please select the "Submit Assignment" button at the top of your lab.