# Programming Assignment 3a

## Assignment Instructions

Write two functions, described below.  Be sure to test your code against the test cases! Submit your code once your tests are passing.  
Note: It is possible to earn partial credit for passing some of the tests, but not passing others.

* #### Function 1:
Define a function that takes no arguments but returns a random number between 1 and 6 (inclusive), using the random package's randrange() function.    
This function can just be 1 line long, depending on how you want to write it. 

* #### Function 2:
Define a second function that takes 3 arguments and returns the sum of the three numbers.  
This function can also be just one line, depending on how you want to write it.

### Extension (optional)

If you want to turn these two functions into a fun program, you could write more lines of code (refer to the live coding videos for how to structure the main parts of your program, and where to place these functions).  This part is totally optional! 

* Start by printing a message telling the user that you are rolling 3 dice. 

* Use the function from above to generate a random number, storing it in a variable, and printing it out. 

* Repeat for a second die. 

* Repeat for a third die. 

* Use the other function from above to add together the values of these 3 dice roles, then print out the total for the user. 

## Let's start coding!

```python
# Install Necessary Dependencies
!pip install nose
import random
from nose.tools import assert_equal
```

Output:
```
Requirement already satisfied: nose in /opt/conda/lib/python3.11/site-packages (1.3.7)
```

### Step 1: Roll a Single dice

Notice that some of the cells are not able to be edited – they are just meant to be run to test your code, and perform the auto grading.  You should begin at the top and execute each code cell, adding your code when you see # YOUR CODE HERE.  If you have written the functions correctly, you should see no errors.

When you execute the TEST CASES cell, what you want to see is no output.  You may want to intentionally make the function give a wrong result, so that you can see what the error message looks like – and then fix your function.  Once you execute the TEST CASES cell, and get no output, then your function is most likely correct!


<!-- grade_id: cell-44a0994fc0d12898, points: 0, grade: false, solution: true -->
```python
### GRADED FUNCTION 1: Write a function below to roll a single 6 sided die.

def roll():
    ''' roll a 6 sided die and return the result as an integer'''
    number = random.randrange(1, 6)
    return number

```

Output:
```
3
```

```python
# Try out a single roll
singleRoll = roll()
print("Your single role value is: ", singleRoll)
```

Output:
```
Your single role value is:  5
```

<!-- grade_id: cell-db80e0ef9e26012a, points: 1, grade: true, solution: false -->
```python
# TEST CASES: This cell should not deliver any errors when this cell is run
assert roll() <= 7 and singleRoll >= 1, "Not quite. Verify that you've accounted for a 6 side die and upper bounds of range. Your roll value %s " % singleRoll
```

### Step 2: Get the total from all dice rolls

<!-- grade_id: cell-58341e91ad0b7cba, points: 0, grade: false, solution: true -->
```python
### GRADED FUNCTION 2:  Write a function called "total_throw" below that adds 3 dice numbers together

def total_throw(die1,die2,die3):
    ''' return the total of 3 dice (6-sided) rolls as an integer'''
    sum_numbers = die1 + die2 + die3
    return sum_numbers

```

```python
# Try out a roll of three dice
totalRoll = total_throw(1, 2, 5)
print("Your throw value is: ", totalRoll)
```

Output:
```
Your throw value is:  8
```

<!-- grade_id: cell-7a38aac5fa573393, points: 1, grade: true, solution: false -->
```python
# GRADED TEST CASES: This cell should not deliver any errors when this cell is run
assert total_throw(1,1,1) <= 18 and totalRoll >= 3 and total_throw(1, 2, 5) == 8, "Not quite. Verify that you've added the total of all dice. Your total roll value is %s " % totalRoll
```

## Putting it all together

### Step 3: Now let's put it all together
Below is a program that will use your functions to roll 3 dice for the user.  You should try running this code several times, and notice how the random numbers are different each time.

```python
### BEGIN OPTIONAL SOLUTION
# Tell the user we're ready to roll
print('Rolling 3 dice...')
# Roll 3 dice
# Get the total of all 3 dice and print
one = roll()
print('You rolled a', one)
two = roll()
print('You rolled a', two)
three = roll()
print('You rolled a', three)
total = total_throw(one, two, three)
# report the total
print('The total of your dice throw is',total )
### END OPTIONAL SOLUTION
```

Output:
```
Rolling 3 dice...
You rolled a 1
You rolled a 3
You rolled a 2
The total of your dice throw is 6
```

```python

```

```python

```